# Exercise 8 - Introduction to Data Visualization with Python

- 🏆 20 points available


---

▶️ Run the code cell below to import `unittest`, a module used for **🧭 Check Your Work** sections and the autograder.


```python
import unittest
import base64
import plotly

tc = unittest.TestCase()
```

---

### 🎯 Exercise 1: Import Packages

#### 👇 Tasks

- ✔️ Import the following Python packages.
  1. `pandas`: Use alias `pd`.
  2. `numpy`: Use alias `np`.
  3. `plotly.express`: Use alias `px`.
  4. `plotly.graph_objects`: Use alias `go`.


```python
# YOUR CODE BEGINS
import pandas as pd
import numpy as np
import plotly.graph_objects as go
import plotly.express as px

# YOUR CODE ENDS
```

#### 🧭 Check Your Work

Run the code cell below to test your solution.

- ✔️ If the code cell runs without errors, you're good to move on.
- ❌ If the code cell produces an error, review your code and fix any mistakes.


<!-- grade_id: exercise-01, points: 2, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE IN THIS CELL
_test_case = "part-01"
_points = 2

tc.assertTrue(
    "pd" in globals(), "Check whether you have correctly imported Pandas with an alias."
)
tc.assertTrue(
    "np" in globals(), "Check whether you have correctly imported NumPy with an alias."
)

print(f"The current plotly version is {plotly.__version__}")
plotly_major_version = int(plotly.__version__.split(".")[0])
tc.assertGreaterEqual(
    plotly_major_version, 5, "Your plotly version should be greater than or equal to 5"
)

tc.assertIsNotNone(
    go.Figure,
    "Check whether you have correctly imported plotly.graph_objects with an alias go.",
)
tc.assertIsNotNone(
    px.scatter,
    "Check whether you have correctly imported plotly.express with an alias px.",
)
```

Output:
```
The current plotly version is 6.3.0
```

---

### 📌 Import dataset

Today, we work with a list of used cars information. The data has been downloaded from [https://www.kaggle.com/datasets/akshaydattatraykhare/car-details-dataset](https://www.kaggle.com/datasets/akshaydattatraykhare/car-details-dataset) without any modification.


▶️ Run the code below to import an cars Dataset. 🚗🚓🚕🛺🚙


```python
# Display all columns
pd.set_option("display.max_columns", 50)

df_cars = pd.read_csv(
    "https://github.com/bdi475/datasets/raw/main/car-dekho-used-cars.csv"
)

display(df_cars)
```

Output:
```
name  year  selling_price  km_driven  \
0                           Maruti 800 AC  2007          60000      70000   
1                Maruti Wagon R LXI Minor  2007         135000      50000   
2                    Hyundai Verna 1.6 SX  2012         600000     100000   
3                  Datsun RediGO T Option  2017         250000      46000   
4                   Honda Amaze VX i-DTEC  2014         450000     141000   
...                                   ...   ...            ...        ...   
4335  Hyundai i20 Magna 1.4 CRDi (Diesel)  2014         409999      80000   
4336           Hyundai i20 Magna 1.4 CRDi  2014         409999      80000   
4337                  Maruti 800 AC BSIII  2009         110000      83000   
4338     Hyundai Creta 1.6 CRDi SX Option  2016         865000      90000   
4339                     Renault KWID RXT  2016         225000      40000   

        fuel seller_type transmission         owner  
0     Petrol  Individual       Manual   First Owner  
1   
```

---

### 🎯 Exercise 2: Number of rows and columns

#### 👇 Tasks

- ✔️ Retrieve the number of rows in `df_cars` to a new variable named `num_rows`.
- ✔️ Retrieve the number of columns in `df_cars` to a new variable named `num_cols`.
- ✔️ Both `num_rows` and `num_cols` should be integer types.

#### 🚀 Hints

- `my_dataframe.shape` returns a tuple containing the number of rows and columns of `my_dataframe`.
- You can retrieve the first element of a tuple using square brackets notation.
  - **Example**: `my_dataframe.shape[0]`


```python
# YOUR CODE BEGINS
num_rows = df_cars.shape[0]
num_cols = df_cars.shape[1]
# YOUR CODE ENDS

print(f"df_cars contains {num_rows} rows and {num_cols} columns.")
```

Output:
```
df_cars contains 4340 rows and 8 columns.
```

#### 🧭 Check Your Work

Run the code cell below to test your solution.

- ✔️ If the code cell runs without errors, you're good to move on.
- ❌ If the code cell produces an error, review your code and fix any mistakes.


<!-- grade_id: exercise-02, points: 2, grade: true, solution: false -->
```python
_test_case = "part-02"
_points = 2

tc.assertEqual(
    num_rows, len(df_cars.index), f"Number of rows should be {len(df_cars.index)}"
)
tc.assertEqual(
    num_cols,
    len(df_cars.columns),
    f"Number of columns should be {len(df_cars.columns)}",
)
```

---

### 🎯 Exercise 3: Selling price box plot (horizontal)

#### 👇 Tasks

- ✔️ Draw a horizontal box plot of `selling_price`.
- ✔️ Store your figure to a variable named `fig`.
- ✔️ Add an appropriate title to your figure.
  - A title should describe your plot (e.g., Salary Box Plot).
- ✔️ Display the figure using `fig.show()`


```python
# YOUR CODE BEGINS
fig = px.box(
    df_cars,
    x="selling_price",
    orientation="h",
    title="Selling Price Distribution (Horizontal)",
)
fig.show()
# YOUR CODE ENDS
```

#### 🔑 Sample output

![Salary distribution horizontal](https://github.com/bdi475/images/blob/main/exercises/intro-to-dataviz/selling_price_box_plot.png?raw=true)


#### 🧭 Check Your Work

Run the code cell below to test your solution.

- ✔️ If the code cell runs without errors, you're good to move on.
- ❌ If the code cell produces an error, review your code and fix any mistakes.


<!-- grade_id: exercise-03, points: 2, grade: true, solution: false -->
```python
_test_case = "part-03"
_points = 2

tc.assertEqual(len(fig.data), 1, "There must be only one plot in your figure")
tc.assertIsNotNone(fig.layout.title.text, "Missing figure title")
tc.assertEqual(fig.data[0].type, "box", "Not a box plot")
tc.assertEqual(
    fig.data[0].orientation, "h", "Your plot should have a horizontal orientation"
)
np.testing.assert_array_equal(fig.data[0].x, df_cars["selling_price"], "Incorrect data")
```

---

### 🎯 Exercise 4: Selling price distribution by number of previous owners

#### 👇 Tasks

- ✔️ Draw horizontal box plots of `selling_price` by `owner`.
- ✔️ Store your figure to a variable named `fig`.
- ✔️ Add an appropriate title to your figure.
- ✔️ Display the figure using `fig.show()`


```python
# YOUR CODE BEGINS
fig = px.box(
    df_cars,
    x="selling_price",
    y="owner",
    orientation="h",
    title="Selling Price Distribution by Number of Previous Owners",
)
fig.show()
# YOUR CODE ENDS
```

#### 🔑 Sample output

![Salary distribution by citizenship status box plots](https://github.com/bdi475/images/blob/main/exercises/intro-to-dataviz/selling_price_by_number_of_owners.png?raw=true)


#### 🧭 Check Your Work

Run the code cell below to test your solution.

- ✔️ If the code cell runs without errors, you're good to move on.
- ❌ If the code cell produces an error, review your code and fix any mistakes.


<!-- grade_id: exercise-04, points: 3, grade: true, solution: false -->
```python
_test_case = "part-04"
_points = 2

tc.assertEqual(len(fig.data), 1, "There must be only one plot in your figure")
tc.assertIsNotNone(fig.layout.title.text, "Missing figure title")
tc.assertEqual(fig.data[0].type, "box", "Not a box plot")
tc.assertEqual(
    fig.data[0].orientation, "h", "Your plot should have a horizontal orientation"
)
np.testing.assert_array_equal(
    fig.data[0].x, df_cars["selling_price"], "Incorrect x-axis data"
)
np.testing.assert_array_equal(fig.data[0].y, df_cars["owner"], "Incorrect y-axis data")
```

---

### 🎯 Exercise 5: Driven distance distribution by transmission

#### 👇 Tasks

- ✔️ Draw horizontal box plots of `km_driven` by `transmission`.
- ✔️ Store your figure to a variable named `fig`.
- ✔️ Add an appropriate title to your figure.
- ✔️ Display the figure using `fig.show()`


```python
# YOUR CODE BEGINS
fig = px.box(
    df_cars,
    x="km_driven",
    y="transmission",
    orientation="h",
    title="Driven distance distribution by transmission",
)
fig.show()
# YOUR CODE ENDS
```

#### 🔑 Sample output

![Salary distribution by performance score box plots](https://github.com/bdi475/images/blob/main/exercises/intro-to-dataviz/distance_by_transmission.png?raw=true)


#### 🧭 Check Your Work

Run the code cell below to test your solution.

- ✔️ If the code cell runs without errors, you're good to move on.
- ❌ If the code cell produces an error, review your code and fix any mistakes.


<!-- grade_id: exercise-05, points: 2, grade: true, solution: false -->
```python
_test_case = "part-05"
_points = 2

tc.assertEqual(len(fig.data), 1, "There must be only one plot in your figure")
tc.assertIsNotNone(fig.layout.title.text, "Missing figure title")
tc.assertEqual(fig.data[0].type, "box", "Not a box plot")
tc.assertEqual(
    fig.data[0].orientation, "h", "Your plot should have a horizontal orientation"
)
np.testing.assert_array_equal(
    fig.data[0].x, df_cars["km_driven"], "Incorrect x-axis data"
)
np.testing.assert_array_equal(
    fig.data[0].y, df_cars["transmission"], "Incorrect y-axis data"
)
```

---

### 🎯 Exercise 6: Selling price distribution by fuel type

#### 👇 Tasks

- ✔️ Draw horizontal box plots of `selling_price` by `fuel`.
- ✔️ Store your figure to a variable named `fig`.
- ✔️ Add an appropriate title to your figure.
- ✔️ Set the height of your figure to `600`.
- ✔️ Display the figure using `fig.show()`

#### 🚀 Hints

```python
fig = px.box(
    my_dataframe,
    x='my_column1',
    y='my_column2',
    title='Plot Title Goes Here',
    height=600
)
fig.show()
```


```python
# YOUR CODE BEGINS
fig = px.box(
    df_cars,
    x="selling_price",
    y="fuel",
    orientation="h",
    title="Selling Price by Fuel Type",
    height=600,
)
fig.show()
# YOUR CODE ENDS
```

#### 🔑 Sample output

![Salary distribution by department box plots](https://github.com/bdi475/images/blob/main/exercises/intro-to-dataviz/selling_price_by_fuel_type.png?raw=true)


#### 🧭 Check Your Work

Run the code cell below to test your solution.

- ✔️ If the code cell runs without errors, you're good to move on.
- ❌ If the code cell produces an error, review your code and fix any mistakes.


<!-- grade_id: exercise-06, points: 2, grade: true, solution: false -->
```python
_test_case = "part-06"
_points = 3

tc.assertEqual(len(fig.data), 1, "There must be only one plot in your figure")
tc.assertIsNotNone(fig.layout.title.text, "Missing figure title")
tc.assertEqual(fig.data[0].type, "box", "Not a box plot")
tc.assertEqual(
    fig.data[0].orientation, "h", "Your plot should have a horizontal orientation"
)
tc.assertEqual(fig.layout.height, 600, "Incorrect height")
np.testing.assert_array_equal(
    fig.data[0].x, df_cars["selling_price"], "Incorrect x-axis data"
)
np.testing.assert_array_equal(fig.data[0].y, df_cars["fuel"], "Incorrect y-axis data")
```

---

### 🎯 Exercise 7: Selling price distribution by seller type

#### 👇 Tasks

- ✔️ Draw horizontal box plots of `selling_price` by `seller_type`.
- ✔️ Store your figure to a variable named `fig`.
- ✔️ Add an appropriate title to your figure.
- ✔️ Set the height of your figure to `700`.
- ✔️ Display the figure using `fig.show()`


```python
# YOUR CODE BEGINS
fig = px.box(
    df_cars,
    x="selling_price",
    y="seller_type",
    orientation="h",
    title="Selling Price Distribution by Seller Type",
    height=700,
)
fig.show()
# YOUR CODE ENDS
```

#### 🔑 Sample output

![Salary distribution by recruitment source box plots](https://github.com/bdi475/images/blob/main/exercises/intro-to-dataviz/selling_price_by_seller_type.png?raw=true)


#### 🧭 Check Your Work

Run the code cell below to test your solution.

- ✔️ If the code cell runs without errors, you're good to move on.
- ❌ If the code cell produces an error, review your code and fix any mistakes.


<!-- grade_id: exercise-07, points: 2, grade: true, solution: false -->
```python
_test_case = "part-07"
_points = 3

tc.assertEqual(len(fig.data), 1, "There must be only one plot in your figure")
tc.assertIsNotNone(fig.layout.title.text, "Missing figure title")
tc.assertEqual(fig.data[0].type, "box", "Not a box plot")
tc.assertEqual(
    fig.data[0].orientation, "h", "Your plot should have a horizontal orientation"
)
tc.assertEqual(fig.layout.height, 700, "Incorrect height")
np.testing.assert_array_equal(
    fig.data[0].x, df_cars["selling_price"], "Incorrect x-axis data"
)
np.testing.assert_array_equal(
    fig.data[0].y, df_cars["seller_type"], "Incorrect y-axis data"
)
```

---

### 🎯 Exercise 8: Selling price histogram

#### 👇 Tasks

- ✔️ Draw a histogram (vertical) of `selling_price` in `df_cars`.
- ✔️ Store your figure to a variable named `fig`.
- ✔️ Add an appropriate title to your figure.
- ✔️ Display the figure using `fig.show()`


```python
# YOUR CODE BEGINS
fig = px.histogram(df_cars, x="selling_price", title="Selling Price Distribution")
fig.show()
# YOUR CODE ENDS
```

#### 🔑 Sample output

![Salary dispersion histogram](https://github.com/bdi475/images/blob/main/exercises/intro-to-dataviz/selling_price_histogram.png?raw=true)


#### 🧭 Check Your Work

Run the code cell below to test your solution.

- ✔️ If the code cell runs without errors, you're good to move on.
- ❌ If the code cell produces an error, review your code and fix any mistakes.


<!-- grade_id: exercise-08, points: 3, grade: true, solution: false -->
```python
_test_case = "part-08"
_points = 2

tc.assertEqual(len(fig.data), 1, "There must be only one plot in your figure")
tc.assertIsNotNone(fig.layout.title.text, "Missing figure title")
tc.assertEqual(fig.data[0].type, "histogram", "Not a histogram")
tc.assertEqual(
    fig.data[0].orientation, "v", "Your plot should have a vertical orientation"
)
np.testing.assert_array_equal(fig.data[0].x, df_cars["selling_price"], "Incorrect data")
```

---

### 🎯 Exercise 9: Distance driven distribution

#### 👇 Tasks

- ✔️ Draw a histogram (vertical) of `km_driven` in `df_cars`.
- ✔️ Store your figure to a variable named `fig`.
- ✔️ Add an appropriate title to your figure.
- ✔️ Display the figure using `fig.show()`


```python
# YOUR CODE BEGINS
fig = px.histogram(df_cars, x="km_driven", title="Distance Driven Distribution")
fig.show()
# YOUR CODE ENDS
```

#### 🔑 Sample output

![Number of absence histogram](https://github.com/bdi475/images/blob/main/exercises/intro-to-dataviz/distance_driven_histogram.png?raw=true)


#### 🧭 Check Your Work

Run the code cell below to test your solution.

- ✔️ If the code cell runs without errors, you're good to move on.
- ❌ If the code cell produces an error, review your code and fix any mistakes.


<!-- grade_id: exercise-09, points: 2, grade: true, solution: false -->
```python
_test_case = "part-09"
_points = 2

tc.assertEqual(len(fig.data), 1, "There must be only one plot in your figure")
tc.assertIsNotNone(fig.layout.title.text, "Missing figure title")
tc.assertEqual(fig.data[0].type, "histogram", "Not a histogram")
tc.assertEqual(
    fig.data[0].orientation, "v", "Your plot should have a vertical orientation"
)
np.testing.assert_array_equal(fig.data[0].x, df_cars["km_driven"], "Incorrect data")
```