# Introduction to Jupyter Notebooks and Pandas

## What is a Jupyter Notebook?

A [Jupyter](https://jupyter.org/) notebook is a document that can contain live code w/ results, visualizations, and rich text. It is widely used in data science and analytics. The cell below is a *code* cell. It contains a block of executable code.

Run the code below by clicking on the cell below and clicking the "Run" button on top (▶).

```python
print(10 + 20)
```

▶️ Run the code cell below to import `unittest`, a module used for **🧭 Check Your Work** sections and the autograder.

```python
import unittest
tc = unittest.TestCase()
```

## Types of cells

There are three different type of cells.

1. Code cell
2. Markdown cell
3. Raw cell

We will most frequently use the first two types of cells.

---

### 🎯 Challenge 1: Find the sum of a list

#### 👇 Tasks

- ✔️ Complete the code cell below to find the sum of all values in `my_list`.
- ✔️ Store the result in a new variable named `result`.

```python
my_list = [11, 20, 52, 91, 90, 75, 74, 20, 21, 10, 14]

# YOUR CODE BEGINS



# YOUR CODE ENDS
print(result)
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix any incorrect parts.

```python
import unittest

tc = unittest.TestCase()

tc.assertEqual(result, 478)
```

---

## Introduction to Pandas

Pandas is a Python *library* for data manipulation and analysis. Although it's used universally in data-related programming applications, it was initially developed for financial analysis by [AQR Capital Management](https://www.aqr.com/).

![Pandas logo](https://github.com/bdi475/notebooks/blob/main/images/pandas-logo.png?raw=true)

Note: A *library* in the context of programming is a collection of functions (and other data) that others have already written for you.

Pandas is popular for many reasons:

1. 🏃🏿‍♀️ It's fast (for most cases where the dataset can be loaded to your memory).
2. 🪒 It supports most of the features required for data manipulation.
3. 💡 Write less code. Get more done.

---

### 🎯 Challenge 2: Import packages

#### 👇 Tasks

- ✔️ Import the following Python packages.
    1. `pandas`: Use alias `pd`.
    2. `numpy`: Use alias `np`.

```python
# YOUR CODE BEGINS

# YOUR CODE ENDS
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

```python
import sys
tc.assertTrue("pd" in globals(), "Check whether you have correctly import Pandas with an alias.")
tc.assertTrue("np" in globals(), "Check whether you have correctly import NumPy with an alias.")
```

---

### It all starts with a `Series`...

The basic building block of Pandas is a `Series`. A `Series` is like a list, but with many more features.

You can create a `Series` by passing a list of values to `pd.Series()`.

```python
s = pd.Series([1, 2, 3, np.nan, 5, 6])

s
```

### Few things to note here

1. These look similar to a Python `list`.
2. The last line of the printed output tells us the data type of values in the `Series` (`dtype: float64`).
- What the heck is `np.nan`?
    - It is used to indicate a "missing value".
    - `np.nan` is NOT the same as `0`.
    
### Differences between a list and a Series

```python
my_list = [1, 2, 3, 4]

print(type(my_list))
display(my_list * 2)
```

```python
my_series = pd.Series([1, 2, 3, 4])

print(type(my_series))
display(my_series * 2)
```

What happens when you multiply a Python `list` by number `2`? It repeats the elements.

How about a `Series`? It multiples each element by `2`!

---

### 🎯 Challenge 3: Create new `Series`

#### 👇 Tasks

- ✔️ Create a new Pandas `Series` named `my_series` with the following three values: `10`, `20`, `30`.

#### 🚀 Hint

The code below creates a new Pandas `Series` with the values `1` and `2`.

```python
my_new_series = pd.Series([1, 2])
```

```python
# YOUR CODE BEGINS
# YOUR CODE ENDS
my_series
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix any incorrect parts.

```python
pd.testing.assert_series_equal(my_series, pd.Series([1, 2, 3]) * 10)
```

---

### Using `Series` methods

A pandas `Series` is similar to a Python `list`. However, a `Series` provides many methods (equivalent to functions) for you to use.

As an example, `num_reviews.mean()` will return the average number of reviews.

```python
reviews_count = [12715, 2274, 2771, 3952, 528, 2766, 724]
num_reviews = pd.Series(reviews_count)

# YOUR CODE HERE
...
```

---

### 🎯 Challenge 4: Create a Pandas DataFrame

#### 👇 Tasks

- ✔️ You are given two lists - `product_names` and `num_reviews` that contain the names of make-up products and the number of reviews on Sephora.com.
- ✔️ Using the two lists, create a new Pandas `DataFrame` named `df_top_products` that has the following two columns:
    1. `product_name`: Names of the products
    2. `num_review`: Number of reviews
- ✔️ Note that the column names are singular.

#### 🚀 Hint

The code below creates a new Pandas `DataFrame` from two series.

```python
my_new_dataframe = pd.DataFrame({
    "column_one": my_series1,
    "column_two": my_series2
})
```

```python
product_names = [
    "Laneige Lip Sleeping Mask",
    "The Ordinary Hyaluronic Acid 2% + B5",
    "Laneige Lip Glowy Balm",
    "Chanel COCO MADEMOISELLE Eau de Parfum"
]

num_reviews = [
    12715,
    2274,
    2766,
    724
]

# YOUR CODE BEGINS



# YOUR CODE ENDS
display(df_top_products)
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix any incorrect parts.

<!-- grade_id: challenge-02, points: 1, grade: true, solution: false -->
```python
pd.testing.assert_frame_equal(
    df_top_products.reset_index(drop=True),
    pd.DataFrame({"product_name": {0: "Laneige Lip Sleeping Mask",
        1: "The Ordinary Hyaluronic Acid 2% + B5",
        2: "Laneige Lip Glowy Balm",
        3: "Chanel COCO MADEMOISELLE Eau de Parfum"},
        "num_review": {0: 12715, 1: 2274, 2: 2766, 3: 724}})
)
```

---

### 📌 Load data

The second part of today's lecture is all about **you**. 👻 Literally.

▶️ Run the code cell below to create a new `DataFrame` named `df_you`.

```python
df_you = pd.read_csv("https://github.com/bdi475/datasets/raw/main/about-you.csv")

# Used to keep a clean copy
df_you_backup = df_you.copy()

# head() displays the first 5 rows of a DataFrame
df_you.head()
```

☝️ **Hold on.** Didn't we always create `DataFrame`s using `pd.DataFrame()`?

Yes. But we can *import* existing data as a Pandas `DataFrame` using `pd.read_csv()`. There are many other similar import methods. For now, we'll mostly use `pd.read_csv()`.

The table below explains each column in `df_you`.

| Column Name             | Description                                               |
|-------------------------|-----------------------------------------------------------|
| name                    | First name                                                |
| major1                  | Major                                                     |
| major2                  | Second major OR minor (blank if no second major or minor) |
| city                    | City the person is from                                   |
| fav_restaurant          | Favorite restaurant (blank if no restaurant was given)    |
| fav_movie               | Favorite movie (blank if no movie was given)              |
| has_iphone              | Whether the person use an iPhone                          |

---

### 📌 Concise summary of a `DataFrame`

👉 A common first step in working with a `DataFrame` is to use the `info()` method. `info()` prints a concise summary of a `DataFrame`.
- Index data type
- Column information: for each column, the following information is displayed:
    - Number of non-missing values
    - Data type of the column
- Memory usage

▶️ Run `df_you.info()` below to see the `info()` method in action.

```python
# YOUR CODE BEGINS
# YOUR CODE ENDS
```

👉 From the result of `df_you.info()`, we can understand a couple of things:

- There are 7 columns.
- Five columns have the `object` data type.
    - In Pandas, a string data type is shown as `object`, not `str`.
        - We will skip the technical discussion for now.
- The second line of the output tells us that there are 40 entries.
- Some columns have 40 non-null values - these columns do not contain any missing value.
- Some columns have less than 40 non-null values - these columns contain one or more missing values.
    - Missing values are displayed as `NaN`.
    - To denote a missing value, use NumPy's `np.nan` (more on this later).

---

### 🎯 Challenge 5: Display first/last/random rows

▶️ Run `df_you.head()` to print the first 5 rows of `df_you`.

```python
# YOUR CODE BEGINS
# YOUR CODE ENDS
```

▶️ Run `df_you.tail(4)` to print the last 4 rows of `df_you`.

```python
# YOUR CODE BEGINS
# YOUR CODE ENDS
```

▶️ Run `df_you.sample(3)` to print 3 randomly sampled rows from `df_you`.

```python
# YOUR CODE BEGINS
# YOUR CODE ENDS
```

<!-- grade_id: challenge-03, points: 1, grade: true, solution: false -->
```python
# Autograder
```

---

### 📌 Number of rows and columns in a `DataFrame`

👉 How many rows and columns does `df_you` have?

▶️ Run `df_you.shape` below to see the *shape* (number of rows and columns) of the database.

```python
# YOUR CODE BEGINS
# YOUR CODE ENDS
```

👉 Can you store the number of rows and columns to variables?

---

- `df_you.shape` returns a `tuple` in `(num_rows, num_cols)` format. 
- What is a `tuple`? 🙀
- A `tuple` is a `list` that cannot be modified once created.

▶️ Run the code cell below to see how a `tuple` is nearly identical to a `list`.

```python
# These two are nearly identical,
# The only difference is that my_tuple cannot be modified once created
my_list = [10, 20]
my_tuple = (10, 20)

print(f"my_list[1]={my_list[1]}")    # prints 20
print(f"my_tuple[1]={my_tuple[1]}")  # also prints 20
```

---

### 🎯 Challenge 6: Find the number of rows and columns in a `DataFrame`

#### 👇 Tasks

- ✔️ Store the number of rows in `df_you` to a new variable named `num_rows`.
- ✔️ Store the number of columns in `df_you` to a new variable named `num_cols`.
- ✔️ Use `.shape`, not `len()`.

```python
# YOUR CODE BEGINS

# YOUR CODE ENDS
print(num_rows)
print(num_cols)
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-04, points: 1, grade: true, solution: false -->
```python
tc.assertEqual(num_rows, len(df_you.index), f"Number of rows should be {len(df_you.index)}")
tc.assertEqual(num_cols, len(df_you.columns), f"Number of columns should be {len(df_you.columns)}")
```