# Pandas Data Types and Column Operations

## Lecture Notes and in-class exercises

▶️ First, run the code cell below to import `unittest`, a module used for **🧭 Check Your Work** sections and the autograder.

```python
import unittest
tc = unittest.TestCase()
```

#### 👇 Tasks

- ✔️ Import the following Python packages.
    1. `pandas`: Use alias `pd`.
    2. `numpy`: Use alias `np`.

```python
### BEGIN SOLUTION
import pandas as pd
import numpy as np
### END SOLUTION
```

#### 🧭 Check your work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

```python
import sys
tc.assertTrue("pd" in globals(), "Check whether you have correctly import Pandas with an alias.")
tc.assertTrue("np" in globals(), "Check whether you have correctly import NumPy with an alias.")
```

---

### 📌 Load data

▶️ Run the code cell below to create a new `DataFrame` named `df_emp`.

```python
# DO NOT CHANGE THE CODE IN THIS CELL
df_emp = pd.DataFrame({
    "emp_id": [30, 40, 10, 20],
    "name": ["Colby", "Adam", "Eli", "Dylan"],
    "dept": ["Sales", "Marketing", "Sales", "Marketing"],
    "office_phone": ["(217)123-4500", np.nan, np.nan, "(217)987-6543"],
    "start_date": ["2017-05-01", "2018-02-01", "2020-08-01", "2019-12-01"],
    "salary": [202000, 185000, 240000, 160500]
})

# Used for intermediate checks
df_emp_backup = df_emp.copy()

df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

---

## 👉 Sorting by Column(s) Recap

You can sort a `DataFrame` using `df.sort_values()`.

![sort_values usage](https://github.com/bdi475/images/blob/main/pandas/sort-values-01.png?raw=true)

---

### 🎯 Challenge 1: Sort by `emp_id` ascending

#### 👇 Tasks

- ✔️ Sort `df_emp` by `emp_id` in **ascending** order.
    - Store the result to a new variable named `df_id_asc`.
- ✔️ `df_emp` should remain unaltered after your code.

#### 🚀 Hints

The code below sorts `my_dataframe` by `some_column` in ascending order and stores the sorted `DataFrame` to a new variable `sorted_dataframe`.

```python
sorted_dataframe = my_dataframe.sort_values("some_column")
```

▶️ Run the code cell below to reset your `df_emp`.

```python
df_emp = df_emp_backup.copy()
df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

```python
### BEGIN SOLUTION
df_id_asc = df_emp.sort_values("emp_id")
### END SOLUTION

df_id_asc
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-01, points: 1, grade: true, solution: false -->
```python
# Challenge 1 Autograder
pd.testing.assert_frame_equal(df_id_asc.reset_index(drop=True),
                              df_emp_backup.sort_values("_".join(["EmP", "iD"]).lower()).reset_index(drop=True))
```

---

### 🎯 Challenge 2: Sort by `emp_id` descending

#### 👇 Tasks

- ✔️ Sort `df_emp` by `emp_id` in **descending** order.
    - Store the result to a new variable named `df_id_desc`.
- ✔️ `df_emp` should remain unaltered after your code.

#### 🚀 Hints

The code below sorts `my_dataframe` by `some_column` in descending order and stores the sorted `DataFrame` to a new variable `sorted_dataframe`.

```python
sorted_dataframe = my_dataframe.sort_values("some_column", ascending=False)
```

▶️ Run the code cell below to reset your `df_emp`.

```python
df_emp = df_emp_backup.copy()
df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

```python
### BEGIN SOLUTION
df_id_desc = df_emp.sort_values("emp_id", ascending=False)
### END SOLUTION

df_id_desc
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
1      40   Adam  Marketing            NaN  2018-02-01  185000
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
2      10    Eli      Sales            NaN  2020-08-01  240000
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-02, points: 1, grade: true, solution: false -->
```python
# Challenge 2 Autograder
pd.testing.assert_frame_equal(df_id_desc.reset_index(drop=True),
                              df_emp_backup.sort_values("_".join(["EmP", "iD"]).lower(), ascending=bool(0)).reset_index(drop=True))
```

---

### 🎯 Challenge 3: Sort by `name` ascending

#### 👇 Tasks

- ✔️ Sort `df_emp` by `name` in **ascending** order.
    - Store the result to a new variable named `df_name_asc`.
- ✔️ `df_emp` should remain unaltered after your code.

#### 🚀 Hints

The code below sorts `my_dataframe` by `some_column` in ascending order and stores the sorted `DataFrame` to a new variable `sorted_dataframe`.

```python
sorted_dataframe = my_dataframe.sort_values("some_column")
```

▶️ Run the code cell below to reset your `df_emp`.

```python
df_emp = df_emp_backup.copy()
df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

```python
### BEGIN SOLUTION
df_name_asc = df_emp.sort_values("name")
### END SOLUTION

df_name_asc
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
1      40   Adam  Marketing            NaN  2018-02-01  185000
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
2      10    Eli      Sales            NaN  2020-08-01  240000
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-03, points: 1, grade: true, solution: false -->
```python
# Challenge 3 Autograder
pd.testing.assert_frame_equal(df_name_asc.reset_index(drop=True),
                              df_emp_backup.sort_values("".join(["nA", "Me"]).lower()).reset_index(drop=True))
```

---

### 🎯 Challenge 4: Sort by `name` descending

#### 👇 Tasks

- ✔️ Sort `df_emp` by `name` in **descending** order.
    - Store the result to a new variable named `df_name_desc`.
- ✔️ `df_emp` should remain unaltered after your code.

#### 🚀 Hints

The code below sorts `my_dataframe` by `some_column` in ascending order and stores the sorted `DataFrame` to a new variable `sorted_dataframe`.

```python
sorted_dataframe = my_dataframe.sort_values("some_column")
```

▶️ Run the code cell below to reset your `df_emp`.

```python
df_emp = df_emp_backup.copy()
df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

```python
### BEGIN SOLUTION
df_name_desc = df_emp.sort_values("name", ascending=False)
### END SOLUTION

df_name_desc
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-04, points: 1, grade: true, solution: false -->
```python
# Challenge 4 Autograder
pd.testing.assert_frame_equal(df_name_desc.reset_index(drop=True),
                              df_emp_backup.sort_values("".join(["nA", "Me"]).lower(), ascending=bool(0)).reset_index(drop=True))
```

---

### 🎯 Challenge 5: Sort by `dept` descending and then by `start_date` ascending

#### 👇 Tasks

- ✔️ Sort `df_emp` by `dept` in **descending** order and then by `start_date` in **ascending** order for people within same departments.
- ✔️ Store the sorted result to a new variable named `df_dept_desc_date_asc`.
- ✔️ `df_emp` should remain unaltered after your code.

#### 🚀 Hints

The code below sorts `my_dataframe` by `some_column` in ascending order and then by `another_column` in descending order for rows with same `same_column` values. It stores the sorted `DataFrame` to a new variable named `sorted_dataframe`.

```python
sorted_dataframe = my_dataframe.sort_values(["some_column", "another_column"], ascending=[True, False])
```

▶️ Run the code cell below to reset your `df_emp`.

```python
df_emp = df_emp_backup.copy()
df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

```python
### BEGIN SOLUTION
df_dept_desc_date_asc = df_emp.sort_values(["dept", "start_date"], ascending=[False, True])
### END SOLUTION

df_dept_desc_date_asc
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
2      10    Eli      Sales            NaN  2020-08-01  240000
1      40   Adam  Marketing            NaN  2018-02-01  185000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-05, points: 1, grade: true, solution: false -->
```python
# Challenge 5 Autograder
pd.testing.assert_frame_equal(df_dept_desc_date_asc.reset_index(drop=True),
                              df_emp_backup.sort_values(["".join(["dE", "Pt"]).lower(), "_".join(["sTarT", "dAtE"]).lower()], ascending=[bool(0), bool(1)]).reset_index(drop=True))
```

---

### 🎯 Challenge 6: Sort by `dept` ascending and then by `salary` descending

#### 👇 Tasks

- ✔️ Sort `df_emp` by `dept` in **ascending** order and then by `salary` in **descending** order.
    - Employees within a same department must be sorted by `salary` in descending order.
    - Store the result to a new variable named `df_dept_asc_salary_desc`.
- ✔️ `df_emp` should remain unaltered after your code.

▶️ Run the code cell below to reset your `df_emp`.

```python
df_emp = df_emp_backup.copy()
df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

```python
### BEGIN SOLUTION
df_dept_asc_salary_desc = df_emp.sort_values(["dept", "salary"], ascending=[True, False])
### END SOLUTION

df_dept_asc_salary_desc
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
1      40   Adam  Marketing            NaN  2018-02-01  185000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
2      10    Eli      Sales            NaN  2020-08-01  240000
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-06, points: 1, grade: true, solution: false -->
```python
# Challenge 6 Autograder
pd.testing.assert_frame_equal(df_dept_asc_salary_desc.reset_index(drop=True),
                              df_emp_backup.sort_values(["".join(["dE", "Pt"]).lower(), "sAlArY".lower()], ascending=[bool(1), bool(0)]).reset_index(drop=True))
```

---

### 🎯 Challenge 7: Sort by `salary` descending in-place

#### 👇 Tasks

- ✔️ Sort `df_emp` by `salary` in **descending** order *in-place*.
    - Directly update `df_emp` without creating a new variable.
    
#### 🚀 Hints

The code below sorts `my_dataframe` by `some_column` in descending order *in-place*.

```python
my_dataframe.sort_values("some_column", ascending=False, inplace=True)
```

▶️ Run the code cell below to reset your `df_emp`.

```python
df_emp = df_emp_backup.copy()
df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

```python
### BEGIN SOLUTION
df_emp.sort_values("salary", ascending=False, inplace=True)
### END SOLUTION

df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
2      10    Eli      Sales            NaN  2020-08-01  240000
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-07, points: 1, grade: true, solution: false -->
```python
# Challenge 7 Autograder
pd.testing.assert_frame_equal(df_emp.reset_index(drop=True),
                              df_emp_backup.sort_values("".join(["sAl", "aRy"]).lower(), ascending=bool(0)).reset_index(drop=True))
```

---

### 🎯 Challenge 8: Sort by `department` and `name` both descending in-place

#### 👇 Tasks

- ✔️ Sort `df_emp` by `dept` and then by `name` for employees in the same department.
    - Sort in **descending** orders for both columns *in-place*.
    - Directly update `df_emp` without creating a new variable.
    
#### 🚀 Hints

The code below sorts `my_dataframe` by `some_column` and `another_column` in descending order *in-place*.

```python
my_dataframe.sort_values(["some_column", "another_column"], ascending=[False, False], inplace=True)
```

▶️ Run the code cell below to reset your `df_emp`.

```python
df_emp = df_emp_backup.copy()
df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

```python
### BEGIN SOLUTION
df_emp.sort_values(["dept", "name"], ascending=[False, False], inplace=True)
### END SOLUTION

df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
2      10    Eli      Sales            NaN  2020-08-01  240000
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
1      40   Adam  Marketing            NaN  2018-02-01  185000
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-08, points: 1, grade: true, solution: false -->
```python
# Challenge 8 Autograder
pd.testing.assert_frame_equal(df_emp.reset_index(drop=True),
                              df_emp_backup.sort_values(["".join(["dE", "Pt"]).lower(), ("nA" + "Me").lower()], ascending=[bool(0), bool(0)]).reset_index(drop=True))
```

---

## 👉 Pandas Data Types

Most of Pandas data types stems from NumPy's dtypes. NumPy supports the following data types:

- `float`
- `int`
- `bool`
- `timedelta64[ns]`
- `datetime64[ns]`

Pandas extend NumPy's type systems. Please refer to the [relevant Pandas documentation page](https://pandas.pydata.org/docs/user_guide/basics.html#dtypes) for the list of all extension types.

You can check the data types of each column using the `dtypes` property.

```python
df_emp.dtypes
```

Output:
```
emp_id           int64
name            object
dept            object
office_phone    object
start_date      object
salary           int64
dtype: object
```

To retrieve the data type of a `Series`, access the `dtype` property.

```python
str(df_emp['name'].dtype)
```

Output:
```
'object'
```

Note that the older versions of Pandas did not have an explicit `string` type and stored strings as `object`s.

Beginning with Pandas `1.0.0`, Pandas supports an explicit type for strings called `StringDtype`.

```python
(df_emp['name'].astype('string')).dtype
```

Output:
```
string[python]
```

---

## 👉 Renaming Column(s)

You can rename a column using `df.rename(columns={"name_before": "name_after"})`. Similar to many Pandas operations, you can perform a rename operation either in-place or out-of-place.

```python
# Rename a column and return a new DataFrame without modifying the original
# df's column names will remain unchanged
df_renamed = df.rename(columns={"name_before": "name_after"})

# Rename a column and update the original DataFrame
df.rename(columns={"name_before": "name_after"}, inplace=True)
```

---

### 🎯 Challenge 9: Rename `office_phone` to `phone_num`

#### 👇 Tasks

- ✔️ Rename `office_phone` column to `phone_num`.
- ✔️ Store the result to a new variable named `df_renamed`.
- ✔️ Your original DataFrame (`df_emp`) should remain unaltered.

#### 🚀 Hints

Use the following code to rename `col_before` column to `col_after` *out-of-place*.

```python
renamed_dataframe = my_dataframe.rename(columns={"col_before": "col_after"})
```

▶️ Run the code cell below to reset your `df_emp`.

```python
df_emp = df_emp_backup.copy()
df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

```python
### BEGIN SOLUTION
df_renamed = df_emp.rename(columns={"office_phone": "phone_num"})
### END SOLUTION

df_renamed
```

Output:
```
emp_id   name       dept      phone_num  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-09, points: 1, grade: true, solution: false -->
```python
# Challenge 9 Autograder
tc.assertEqual(df_emp.columns.tolist(), df_emp_backup.columns.tolist(), "Did you rename the column in-place? The original DataFrame should not be modified.")
tc.assertEqual(df_renamed.columns.tolist(), ["emp_id", "name", "dept", "phone_num", "start_date", "salary"])
```

---

### 🎯 Challenge 10: Rename `office_phone` to `phone_num` **in-place**

#### 👇 Tasks

- ✔️ Rename `office_phone` column to `phone_num` *in-place*.
    - Directly update `df_emp` without creating a new variable.
    
#### 🚀 Hints

Use the following code to rename `col_before` column to `col_after` *in-place*.

```python
my_dataframe.rename(columns={"col_before": "col_after"}, inplace=True)
```

▶️ Run the code cell below to reset your `df_emp`.

```python
df_emp = df_emp_backup.copy()
df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

```python
### BEGIN SOLUTION
df_emp.rename(columns={"office_phone": "phone_num"}, inplace=True)
### END SOLUTION

df_emp
```

Output:
```
emp_id   name       dept      phone_num  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-10, points: 1, grade: true, solution: false -->
```python
# Challenge 10 Autograder
tc.assertEqual(df_emp.columns.tolist(), ["emp_id", "name", "dept", "phone_num", "start_date", "salary"])
```

---

### 🎯 Challenge 11: Rename `name` to `first_name` and `salary` to `base_salary` **in-place**

#### 👇 Tasks

- ✔️ Rename `name` column to `first_name` and `salary` to `base_salary` *in-place*.
    - Directly update `df_emp` without creating a new variable.
    
#### 🚀 Hints

Use the following code as a reference.

```python
my_dataframe.rename(columns={"col_before1": "col_after1", "col_before2": "col_after2"}, inplace=True)
```

▶️ Run the code cell below to reset your `df_emp`.

```python
df_emp = df_emp_backup.copy()
df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

```python
### BEGIN SOLUTION
df_emp.rename(columns={"name": "first_name", "salary": "base_salary"}, inplace=True)
### END SOLUTION

df_emp
```

Output:
```
emp_id first_name       dept   office_phone  start_date  base_salary
0      30      Colby      Sales  (217)123-4500  2017-05-01       202000
1      40       Adam  Marketing            NaN  2018-02-01       185000
2      10        Eli      Sales            NaN  2020-08-01       240000
3      20      Dylan  Marketing  (217)987-6543  2019-12-01       160500
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-11, points: 1, grade: true, solution: false -->
```python
# Challenge 11 Autograder
tc.assertEqual(df_emp.columns.tolist(), ["emp_id", "first_name", "dept", "office_phone", "start_date", "base_salary"])
```

---

## 👉 Dropping Column(s)

You can rename a column using `df.drop(columns=["col1", "col2"])`. Again, you can perform this operation either in-place or out-of-place.

```python
# Copy df, drop columns
# df will remain unchaged
df_dropped = df.drop(columns=["col1", "col2"])

# Drop columns directly from df
df.drop(columns=["col1", "col2"], inplace=True)
```

---

### 🎯 Challenge 12: Drop `start_date` column

#### 👇 Tasks

- ✔️ Drop `start_date` column from `df_emp`.
- ✔️ Store the result to a new variable named `df_dropped`.
- ✔️ Your `df_emp` should remain unaltered.

#### 🚀 Hints

Use the following code as a reference.

```python
dropped_dataframe = my_dataframe.drop(columns=["my_column1"])
```

▶️ Run the code cell below to reset your `df_emp`.

```python
df_emp = df_emp_backup.copy()
df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

```python
### BEGIN SOLUTION
df_dropped = df_emp.drop(columns=["start_date"])
### END SOLUTION

df_dropped
```

Output:
```
emp_id   name       dept   office_phone  salary
0      30  Colby      Sales  (217)123-4500  202000
1      40   Adam  Marketing            NaN  185000
2      10    Eli      Sales            NaN  240000
3      20  Dylan  Marketing  (217)987-6543  160500
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-12, points: 1, grade: true, solution: false -->
```python
# Challenge 12 Autograder
tc.assertEqual(df_emp.columns.tolist(), df_emp_backup.columns.tolist())
tc.assertEqual(df_dropped.columns.tolist(), ["emp_id", "name", "dept", "office_phone", "salary"])
```

---

### 🎯 Challenge 13: Drop `start_date` column **in-place**

#### 👇 Tasks

- ✔️ Drop `start_date` column *in-place*.
    - Directly update `df_emp` without creating a new variable.
    
#### 🚀 Hints

Use the following code as a reference.

```python
my_dataframe.drop(columns=["my_column1"], inplace=True)
```

▶️ Run the code cell below to reset your `df_emp`.

```python
df_emp = df_emp_backup.copy()
df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

```python
### BEGIN SOLUTION
df_emp.drop(columns=["start_date"], inplace=True)
### END SOLUTION

df_emp
```

Output:
```
emp_id   name       dept   office_phone  salary
0      30  Colby      Sales  (217)123-4500  202000
1      40   Adam  Marketing            NaN  185000
2      10    Eli      Sales            NaN  240000
3      20  Dylan  Marketing  (217)987-6543  160500
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-13, points: 1, grade: true, solution: false -->
```python
# Challenge 13 Autograder
tc.assertEqual(df_emp.columns.tolist(), ["emp_id", "name", "dept", "office_phone", "salary"])
```

---

### 🎯 Challenge 14: Drop `name` and `salary` columns **in-place**

#### 👇 Tasks

- ✔️ Drop `name` and `salary` columns *in-place*.
    - Directly update `df_emp` without creating a new variable.
    
#### 🚀 Hints

Use the following code as a reference.

```python
my_dataframe.drop(columns=["my_column1", "my_column2"], inplace=True)
```

▶️ Run the code cell below to reset your `df_emp`.

```python
df_emp = df_emp_backup.copy()
df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

```python
### BEGIN SOLUTION
df_emp.drop(columns=["name", "salary"], inplace=True)
### END SOLUTION

df_emp
```

Output:
```
emp_id       dept   office_phone  start_date
0      30      Sales  (217)123-4500  2017-05-01
1      40  Marketing            NaN  2018-02-01
2      10      Sales            NaN  2020-08-01
3      20  Marketing  (217)987-6543  2019-12-01
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-14, points: 1, grade: true, solution: false -->
```python
# Challenge 14 Autograder
tc.assertEqual(df_emp.columns.tolist(), ["emp_id", "dept", "office_phone", "start_date"])
```

---

## 👉 Selecting Subset Columns

You can select a subset of columns from a `DataFrame` using `df[list_of_columns]`.

You can also use `df.loc[:, list_of_columns]`.

```python
df[["col1", "col2"]]

# is equivalent to

df.loc[:, ["col1", "col2"]]
```

---

### 🎯 Challenge 15: `emp_id` and `name`

#### 👇 Tasks

- ✔️ Select only `emp_id` and `name` columns from `df_emp` (in the same order).
- ✔️ Store the result to a new variable named `df_id_name`.
- ✔️ Your `df_emp` should remain unaltered.

#### 🚀 Hints

Use the following code as a reference.

```python
df_subset = df[["col1", "col2"]]
```

▶️ Run the code cell below to reset your `df_emp`.

```python
df_emp = df_emp_backup.copy()
df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

```python
### BEGIN SOLUTION
df_id_name = df_emp[["emp_id", "name"]]
### END SOLUTION

df_id_name
```

Output:
```
emp_id   name
0      30  Colby
1      40   Adam
2      10    Eli
3      20  Dylan
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-15, points: 1, grade: true, solution: false -->
```python
# Challenge 15 Autograder
pd.testing.assert_frame_equal(df_emp, df_emp_backup)
pd.testing.assert_frame_equal(df_emp_backup[["emp_id", "name"]], df_id_name, check_like=True)
```

---

### 🎯 Challenge 16: `emp_id`, `name`, `dept`, `salary`

#### 👇 Tasks

- ✔️ Select `emp_id`, `name`, `dept`, and `salary` columns from `df_emp` (in the same order).
- ✔️ Store the result to a new variable named `df_subset`.
- ✔️ Your `df_emp` should remain unaltered.

#### 🚀 Hints

Use the following code as a reference.

```python
df_subset = df[["col1", "col2"]]
```

▶️ Run the code cell below to reset your `df_emp`.

```python
df_emp = df_emp_backup.copy()
df_emp
```

Output:
```
emp_id   name       dept   office_phone  start_date  salary
0      30  Colby      Sales  (217)123-4500  2017-05-01  202000
1      40   Adam  Marketing            NaN  2018-02-01  185000
2      10    Eli      Sales            NaN  2020-08-01  240000
3      20  Dylan  Marketing  (217)987-6543  2019-12-01  160500
```

```python
### BEGIN SOLUTION
df_subset = df_emp[["emp_id", "name", "dept", "salary"]]
### END SOLUTION

df_subset
```

Output:
```
emp_id   name       dept  salary
0      30  Colby      Sales  202000
1      40   Adam  Marketing  185000
2      10    Eli      Sales  240000
3      20  Dylan  Marketing  160500
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-16, points: 1, grade: true, solution: false -->
```python
# Challenge 16 Autograder
pd.testing.assert_frame_equal(df_emp, df_emp_backup)
pd.testing.assert_frame_equal(df_emp_backup[["emp_id", "name", "dept", "salary"]], df_subset)
```

---

## 👉 Working with Missing Values

A dataset commonly can contain one or more missing values. Pandas provides a flexible set of methods for you to work with missing values. Our first discussion of missing values will demonstrate how to find missing values.

Missing values can be denoted by multiple *markers*. Here are a few common markers.

- `NaN`: NaN stands for "not a number". For simplicity and performance's sake, Pandas has traditionally used `NaN` to denote missing values.
- `None` (python's built-in type)
- `NA`
- `NaT`

---

### 📌 Load data

▶️ Run the code cell below to create a new `DataFrame` named `df_you`.

```python
# DO NOT CHANGE THE CODE IN THIS CELL
df_you = pd.read_csv("https://github.com/bdi475/datasets/raw/main/about-you.csv")

# Used to keep a clean copy
df_you_backup = df_you.copy()

# head() displays the first 5 rows of a DataFrame
df_you.head()
```

Output:
```
name                major1  \
0  Ahana Chakraborty            Statistics   
1      Andrew Rozmus            Psychology   
2       Anusha Adira  Computer Engineering   
3     Arthur Pyptyuk             Economics   
4        Aryajit Das             Economics   

                                   major2             city country  \
0                  Business & Informatics          Chicago     USA   
1                                     NaN         Elmhurst     USA   
2                                Business        Cupertino     USA   
3                                Business  Hoffman Estates     USA   
4  Business & Global Markets plus Society       Streamwood     USA   

  fav_restaurant                        fav_movie  has_iphone  
0       Poke Lab                          Shrek 2        True  
1            NaN                              NaN        True  
2   Bangkok Thai                     Three Idiots        True  
3       Sakanaya                       Hereditary        True 
```

The table below describes each column in `df_you`.

| Column Name             | Description                                               |
|-------------------------|-----------------------------------------------------------|
| name                    | First name                                                |
| major1                  | Major                                                     |
| major2                  | Second major OR minor (blank if no second major or minor) |
| city                    | City the person is from                                   |
| distance_from_champaign | Straight distance from the city to Champaign in miles     |
| fav_restaurant          | Favorite restaurant (blank if no restaurant was given)    |
| has_iphone              | Whether the person use an iPhone                          |

---

### 🎯 Challenge 17: People from Naperville

#### 👇 Tasks

- ✔️ Using `df_you`, filter rows where the person is from `"Naperville"`.
    - Check whether the `city` column is equal to `"Naperville"`.
    - Store the result to a new variable named `df_naperville`.
- ✔️ `df_you` should remain unaltered after your code.

```python
### BEGIN SOLUTION
df_naperville = df_you[df_you["city"] == "Naperville"]
### END SOLUTION

df_naperville
```

Output:
```
name                  major1    major2        city country  \
9           Dev Singh        Computer Science  Business  Naperville     USA   
23  Kriti Gangadharan  Industrial Engineering  Business  Naperville     USA   

   fav_restaurant      fav_movie  has_iphone  
9     Silvercreek            NaN        True  
23    Sticky Rice  The Holdovers        True
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-17, points: 1, grade: true, solution: false -->
```python
# Challenge 17 Autograder
# df_you should remain unaltered
pd.testing.assert_frame_equal(df_you, df_you_backup)

pd.testing.assert_frame_equal(df_naperville.sort_values(df_naperville.columns.tolist()).reset_index(drop=True),
                              df_you_backup.query(f'{"cItY".lower()} == "{"NaPerVillE".capitalize()}"')
                                 .sort_values(df_you_backup.columns.tolist()).reset_index(drop=True))
```

---

### 🎯 Challenge 18: Anyone with a non-missing `major2`

#### 👇 Tasks

- ✔️ Using `df_you`, filter rows where the person has a second major or a minor.
    - You're looking for rows where `major2` is not `NaN`.
- ✔️ `NaN` is a special value to denote missing value. You must use `my_series.isna()` or `my_series.notna()` to check whether a row contains a missing value or not.
- ✔️ Store the result to a new variable named `df_major2`.
- ✔️ `df_you` should remain unaltered after your code.

#### 🚀 Hints

- `my_series.notna()` can be used to check whether a row contains a missing value or not.

![notna](https://github.com/bdi475/images/blob/main/pandas/notna-series.png?raw=true)

```python
### BEGIN SOLUTION
df_major2 = df_you[df_you["major2"].notna()]
### END SOLUTION

df_major2
```

Output:
```
name                     major1  \
0    Ahana Chakraborty                 Statistics   
2         Anusha Adira       Computer Engineering   
3       Arthur Pyptyuk                  Economics   
4          Aryajit Das                  Economics   
5        Avery Ahrendt                  Economics   
6          Caroline Gu              Communication   
8          Cole Jordan           Computer Science   
9            Dev Singh           Computer Science   
10         Dianne Park  Information Sciences + DS   
12       Ella Crockett          Industrial Design   
15        Halle Maffia                Advertising   
16        Hannah Tabor                Advertising   
17          Hugh Palin       Computer Engineering   
18       Indigo Anstee                 Psychology   
19           Jenny Hou   Interdisciplinary Health   
21         Julia Kevin             Bioengineering   
22       Katie Holland  Earth, Soc, Env Sustainab   
23   Kriti Gangadharan     Industrial Engineering   
25        L
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-18, points: 1, grade: true, solution: false -->
```python
# Challenge 18 Autograder
# df_you should remain unaltered
pd.testing.assert_frame_equal(df_you, df_you_backup)

pd.testing.assert_frame_equal(df_major2.sort_values(df_major2.columns.tolist()).reset_index(drop=True),
                              df_you_backup.query(f"major2 == major2")
                                 .sort_values(df_you_backup.columns.tolist()).reset_index(drop=True))
```

---

### 🎯 Challenge 19: Anyone without a favorite movie

#### 👇 Tasks

- ✔️ Using `df_you`, filter rows where the person's `fav_movie` is `NaN`.
- ✔️ `NaN` is a special value to denote missing value. You must use `my_series.isna()` or `my_series.notna()` to compare a `Series` with `NaN`.
- ✔️ Store the result to a new variable named `df_no_fav_movie`.
- ✔️ `df_you` should remain unaltered after your code.

#### 🚀 Hints

- `my_series.isna()` can be used to check whether a row contains a missing value.

```python
### BEGIN SOLUTION
df_no_fav_movie = df_you[df_you["fav_movie"].isna()]
### END SOLUTION

df_no_fav_movie
```

Output:
```
name                     major1  \
1         Andrew Rozmus                 Psychology   
7   Charan Govarthanara                 Statistics   
9             Dev Singh           Computer Science   
11             Eben Lee     Mechanical Engineering   
14       Gauthami Yenne       Computer Engineering   
20       Jeremy Estrada     Industrial Engineering   
24             Kyler Yu           Computer Science   
26             Litao Li           Computer Science   
29        Nelson Oshodi                        NaN   
30        Nieve Boland              Bioengineering   
33       Paul Muehlberg  Econometrics & Quant Econ   
34         Ridah Shaikh       Information Sciences   
35           Runjun Lan              Communication   
40               Tom Fu                  Economics   

                                      major2              city country  \
1                                        NaN          Elmhurst     USA   
7                                        NaN               N
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-19, points: 1, grade: true, solution: false -->
```python
# Challenge 19 Autograder
# df_you should remain unaltered
pd.testing.assert_frame_equal(df_you, df_you_backup)

pd.testing.assert_frame_equal(df_no_fav_movie.sort_values(df_no_fav_movie.columns.tolist()).reset_index(drop=True),
                              df_you_backup.query(f"fav_movie != fav_movie")
                                 .sort_values(df_you_backup.columns.tolist()).reset_index(drop=True))
```