# Sentiment Analysis with Vader

## 🧭 Overview & Setup

This notebook walks you through on:

1. How to use the [VADER](https://github.com/cjhutto/vaderSentiment) package to analyze sentiments
2. How to run a sentiment analysis on tweets using the [VADER](https://github.com/cjhutto/vaderSentiment) package
3. How to interpret the VADER scores
4. How to view top positive/negative tweets
5. How to visualize the compound scores distribution

▶️ First, run the code cell below to import `unittest`, a module used for **🧭 Check Your Work** sections and the autograder.

```python
import unittest
tc = unittest.TestCase()
```

---

### 🎯 Challenge 1: Import packages

#### 👇 Tasks

- ✔️ Import the following Python packages.
    1. `pandas`: Use alias `pd`.
    2. `numpy`: Use alias `np`.

```python
### BEGIN SOLUTION
import pandas as pd
import numpy as np
### END SOLUTION
```

#### 🧭 Check your work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

```python
# DO NOT CHANGE THE CODE BELOW
_test_case = 'import-packages'
_points = 1

import sys
tc.assertTrue("pd" in globals(), "Check whether you have correctly import Pandas with an alias.")
tc.assertTrue("np" in globals(), "Check whether you have correctly import NumPy with an alias.")
```

---

### 📌 Install and load the `vaderSentiment` package

▶️ Run the code cell below to install `vaderSentiment`.

```python
# install vaderSentiment
!pip install vaderSentiment
```

Output:
```
Defaulting to user installation because normal site-packages is not writeable
Requirement already satisfied: vaderSentiment in /home/jupyter-subwaymatch/.local/lib/python3.9/site-packages (3.3.2)
Requirement already satisfied: requests in /opt/tljh/user/lib/python3.9/site-packages (from vaderSentiment) (2.28.1)
Requirement already satisfied: urllib3<1.27,>=1.21.1 in /opt/tljh/user/lib/python3.9/site-packages (from requests->vaderSentiment) (1.26.7)
Requirement already satisfied: certifi>=2017.4.17 in /opt/tljh/user/lib/python3.9/site-packages (from requests->vaderSentiment) (2022.12.7)
Requirement already satisfied: charset-normalizer<3,>=2 in /opt/tljh/user/lib/python3.9/site-packages (from requests->vaderSentiment) (2.0.0)
Requirement already satisfied: idna<4,>=2.5 in /opt/tljh/user/lib/python3.9/site-packages (from requests->vaderSentiment) (3.1)
```

▶️ Run the code cell below to initialize `vaderSentiment.SentimentIntensityAnalyzer`.

```python
from vaderSentiment.vaderSentiment import SentimentIntensityAnalyzer

# create an analyzer instance
analyzer = SentimentIntensityAnalyzer()
```

▶️ Run the code cell below to display full tweet text in the notebook outputs.

```python
pd.set_option('display.max_colwidth', 1000)
```

---

## 🔮 Sentiment Analysis

### 📌 Using the `analyzer.polarity_scores()` method

To analyze a piece of text using VADER, use the `analyzer.polarity_scores()` method. Here is an example using a string.

compound score is the overall score between -1 (most extreme negative) and +1 (most extreme positive).
pos, neg, and neu are ratios for proportions of text that fall in each category. These should all add up to be 1.
We are mainly interested in the compound score.

```python
vader_score = analyzer.polarity_scores("Got my BeautyBase merch I am happy I got the first batch 🥰")
print(vader_score)
```

```python
vader_score = analyzer.polarity_scores("Got my BeautyBase merch I am happy I got the first batch 🥰")
print(vader_score)
```

Output:
```
{'neg': 0.0, 'neu': 0.56, 'pos': 0.44, 'compound': 0.9001}
```

Here is another example.

```python
sentence = "Beautyrite same day shipping is phenomenal its unbelievable i got my package in few hours"

vader_score = analyzer.polarity_scores(sentence)
print(vader_score)
```

Output:
```
{'neg': 0.0, 'neu': 0.886, 'pos': 0.114, 'compound': 0.2023}
```

---

### 🎯 Challenge 2: Calculate vader score of a string

#### 👇 Tasks

- ✔️ Calculate the polarity scores of the string below.
- ✔️ Store the result to a new variable named `vader_score1`.

```plaintext
I have never been spoken to so terrible and lied to so many times from what should be professionals. I was trying to help my daughter rent a home and was lied to numerous times about their fair housing practices and was given different information from different individuals. This seems to be a very poorly ran company that I would NEVER trust to oversee my daughters safe housing experience! As a parent do not ever allow your adult child to rent from this company! They will not be taken care of!
```

#### 🔑 Output Format

```
{'neg': 0.162, 'neu': 0.765, 'pos': 0.074, 'compound': -0.8565}
```

```python
### BEGIN SOLUTION
s = "I have never been spoken to so terrible and lied to so many times from what should be professionals. I was trying to help my daughter rent a home and was lied to numerous times about their fair housing practices and was given different information from different individuals. This seems to be a very poorly ran company that I would NEVER trust to oversee my daughters safe housing experience! As a parent do not ever allow your adult child to rent from this company! They will not be taken care of!"

vader_score1 = analyzer.polarity_scores(s)
### END SOLUTION

print(vader_score1)
```

Output:
```
{'neg': 0.162, 'neu': 0.765, 'pos': 0.074, 'compound': -0.8565}
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-03, points: 1, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE BELOW
_test_case = 'calculate-vader-score-of-a-string'
_points = 1
_obfuscate = True

s_test = "I have never been spoken to so terrible and lied to so many times from what should be professionals. I was trying to help my daughter rent a home and was lied to numerous times about their fair housing practices and was given different information from different individuals. This seems to be a very poorly ran company that I would NEVER trust to oversee my daughters safe housing experience! As a parent do not ever allow your adult child to rent from this company! They will not be taken care of!"

vader_score1_SOL = analyzer.polarity_scores(s)

tc.assertEqual(vader_score1, vader_score1_SOL)
```

---

### 📌 Applying `analyzer.polarity_scores()` on a DataFrame column.

To apply VADER to a pandas Series of strings, combine `Series.apply()` and `analyzer.polarity_scores()` methods. Below is an example to find the polarity scores on a DataFrame.

▶️ Run the code cell below to create a DataFrame with one string column.

```python
sample_text_values = [
    "VADER is smart, handsome, and funny.",  # positive sentence example
    "VADER is smart, handsome, and funny!",  # punctuation emphasis handled correctly (sentiment intensity adjusted)
    "VADER is very smart, handsome, and funny.", # booster words handled correctly (sentiment intensity adjusted)
    "VADER is VERY SMART, handsome, and FUNNY.",  # emphasis for ALLCAPS handled
    "VADER is VERY SMART, handsome, and FUNNY!!!", # combination of signals - VADER appropriately adjusts intensity
    "VADER is VERY SMART, uber handsome, and FRIGGIN FUNNY!!!", # booster words & punctuation make this close to ceiling for score
    "VADER is not smart, handsome, nor funny.",  # negation sentence example
    "The book was good.",  # positive sentence
    "At least it isn't a horrible book.",  # negated negative sentence with contraction
    "The book was only kind of good.", # qualified positive sentence is handled correctly (intensity adjusted)
    "The plot was good, but the characters are uncompelling and the dialog is not great.", # mixed negation sentence
    "Today SUX!",  # negative slang with capitalization emphasis
    "Today only kinda sux! But I'll get by, lol", # mixed sentiment example with slang and constrastive conjunction "but"
    "Make sure you :) or :D today!",  # emoticons handled
    "Catch utf-8 emoji such as such as 💘 and 💋 and 😁",  # emojis handled
    "Not bad at all"  # Capitalized negation
]

df_sample = pd.DataFrame({
    'text': sample_text_values
})

df_sample
```

Output:
```
text
0                                                  VADER is smart, handsome, and funny.
1                                                  VADER is smart, handsome, and funny!
2                                             VADER is very smart, handsome, and funny.
3                                             VADER is VERY SMART, handsome, and FUNNY.
4                                           VADER is VERY SMART, handsome, and FUNNY!!!
5                              VADER is VERY SMART, uber handsome, and FRIGGIN FUNNY!!!
6                                              VADER is not smart, handsome, nor funny.
7                                                                    The book was good.
8                                                    At least it isn't a horrible book.
9                                                       The book was only kind of good.
10  The plot was good, but the characters are uncompelling and the dialog is not great.
11                         
```

▶️ Run the code cell below to calculate the polarity scores for each row and convert the output into a DataFrame.

```python
# for each row of the text column, run the polarity_scores() function
vader_scores = df_sample['text'].apply(analyzer.polarity_scores)

# convert the Series of dictionaries into a DataFrame
df_vader_scores = pd.DataFrame(
    vader_scores.tolist()
)

df_vader_scores.head()
```

Output:
```
neg    neu    pos  compound
0  0.0  0.254  0.746    0.8316
1  0.0  0.248  0.752    0.8439
2  0.0  0.299  0.701    0.8545
3  0.0  0.246  0.754    0.9227
4  0.0  0.233  0.767    0.9342
```

▶️ Run the code cell below to concatenate `df_sample` and `df_vader_scores` horizontally.

```python
df_merged = pd.concat([df_sample, df_vader_scores], axis=1)
df_merged
```

Output:
```
text  \
0                                                  VADER is smart, handsome, and funny.   
1                                                  VADER is smart, handsome, and funny!   
2                                             VADER is very smart, handsome, and funny.   
3                                             VADER is VERY SMART, handsome, and FUNNY.   
4                                           VADER is VERY SMART, handsome, and FUNNY!!!   
5                              VADER is VERY SMART, uber handsome, and FRIGGIN FUNNY!!!   
6                                              VADER is not smart, handsome, nor funny.   
7                                                                    The book was good.   
8                                                    At least it isn't a horrible book.   
9                                                       The book was only kind of good.   
10  The plot was good, but the characters are uncompelling and the dialog is not g
```

---

### 📌 Load tweets

▶️ Run the code cell below to create a new `DataFrame` containing a list of tweets.

```python
df_tweets = pd.read_csv('https://github.com/bdi475/datasets/raw/main/fake-tweets-beautyrite.csv')
df_tweets_backup = df_tweets.copy()
df_tweets.head(10)
```

Output:
```
text
0                                                                                            every time I go into Beautyrite everybody’s makeup look bad, what’s that about?
1                                                                                                               Got in my car to go to the gym but it took me to Beautyrite.
2                                                                     Why don’t airports have helpful stores like Beautyrite? Why do we need 5 different burger restaurants?
3                                                                                                   Y’all go to Beautyrite’s IG and look at their awful spring lineup 😭😭😭😭😭😭
4                                                  Today I learned that Acme Makeups, BeautyRite, Wonka Factory are all owned by a single company. The illusion of choice...
5                      My kid has entered the “how many chemicals and colors can I put on my face in one day” phase. Please help, 
```

---

### 🎯 Challenge 3: Calculate vader score of a text column

#### 👇 Tasks

- ✔️ Using `df_tweets`, calcualte the polarity scores of each row.
- ✔️ Concatenate the polarity scores horizontally to `df_tweets`.
    - Store the concatenated DataFrame to a new variable named `df_merged`.

#### 🔑 Expected Output

|    | text                                                                                                                      |   neg |   neu |   pos |   compound |
|---:|:--------------------------------------------------------------------------------------------------------------------------|------:|------:|------:|-----------:|
|  0 | every time I go into Beautyrite everybody’s makeup look bad, what’s that about?                                           | 0.226 | 0.774 | 0.000 |    -0.5423 |
|  1 | Got in my car to go to the gym but it took me to Beautyrite.                                                              | 0.000 | 1.000 | 0.000 |    0.0000  |
|  2 | Why don’t airports have helpful stores like Beautyrite? Why do we need 5 different burger restaurants?                    | 0.000 | 0.712 | 0.288 |     0.6868 |
|  3 | Y’all go to Beautyrite’s IG and look at their awful spring lineup 😭😭😭😭😭😭                                            | 0.484 | 0.516 | 0.000 |    -0.9666 |
|  4 | Today I learned that Acme Makeups, BeautyRite, Wonka Factory are all owned by a single company. The illusion of choice... | 0.000 | 1.000 | 0.000 |    0.0000  |

```python
### BEGIN SOLUTION
vader_scores = df_tweets['text'].apply(analyzer.polarity_scores)

df_vader_scores = pd.DataFrame(
    vader_scores.tolist()
)

df_merged = pd.concat([df_tweets, df_vader_scores], axis=1)
### END SOLUTION

df_merged.head(5)
```

Output:
```
text  \
0                                            every time I go into Beautyrite everybody’s makeup look bad, what’s that about?   
1                                                               Got in my car to go to the gym but it took me to Beautyrite.   
2                     Why don’t airports have helpful stores like Beautyrite? Why do we need 5 different burger restaurants?   
3                                                   Y’all go to Beautyrite’s IG and look at their awful spring lineup 😭😭😭😭😭😭   
4  Today I learned that Acme Makeups, BeautyRite, Wonka Factory are all owned by a single company. The illusion of choice...   

     neg    neu    pos  compound  
0  0.226  0.774  0.000   -0.5423  
1  0.000  1.000  0.000    0.0000  
2  0.000  0.712  0.288    0.6868  
3  0.484  0.516  0.000   -0.9666  
4  0.000  1.000  0.000    0.0000
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-03, points: 1, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE BELOW
_test_case = 'calculate-vader-scores-of-a-text-column'
_points = 1
_obfuscate = True

vader_scores_SOL = df_tweets_backup['text'].apply(analyzer.polarity_scores)

df_vader_scores_SOL = pd.DataFrame(
    vader_scores_SOL.tolist()
)

df_merged_SOL = pd.concat([df_tweets_backup, df_vader_scores_SOL], axis=1)
pd.testing.assert_frame_equal(df_merged.sort_values(df_merged.columns.to_list()).reset_index(drop=True),
                              df_merged_SOL.sort_values(df_merged_SOL.columns.tolist()).reset_index(drop=True))
```

---

### 🎯 Challenge 4: Find the average compound score

#### 👇 Tasks

- ✔️ Using `df_merged`, calcualte the average compound score.
- ✔️ Store the result to a new variable named `average_compound_score`.

```python
### BEGIN SOLUTION
average_compound_score = df_merged['compound'].mean()
### END SOLUTION

print(f"The average compound score is {round(average_compound_score, 4)}")
```

Output:
```
The average compound score is 0.1124
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-03, points: 1, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE BELOW
_test_case = 'calculate-average-compound-score'
_points = 1
_obfuscate = True

average_compound_score_SOL = df_merged_SOL['compound'].mean()

tc.assertAlmostEqual(average_compound_score, average_compound_score_SOL)
```

---

### 🎯 Challenge 5: Find the top 10 positive tweets

#### 👇 Tasks

- ✔️ Using `df_merged`, filter the top 10 postive tweets.
    - Use the compound score (`compound` column).
- ✔️ Store the filtered DataFrame to a new variable named `df_top10_positive`.
- ✔️ Sort `df_top10_positive` by `compound` in descending order.

#### 🔑 Expected Output

|    | text                                                                                                                                                                                                                 |   neg |   neu |   pos |   compound |
|---:|:---------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------|------:|------:|------:|-----------:|
| 37 | My BFF gave me a Beautyrite gift card for my birthday 🥰 i love getting gift cards.                                                                                                                                  | 0.000 | 0.406 | 0.594 |     0.9717 |
| 24 | IN LOVE WITH BEAUTY2LIPS 🤩🥰                                                                                                                                                                                        | 0.000 | 0.364 | 0.636 |     0.9222 |
| 19 | Got my BeautyBase merch I am happy I got the first batch 🥰                                                                                                                                                          | 0.000 | 0.56  | 0.44  |     0.9001 |
|  6 | Beautyrite is 20% off today (all purchases). 30% off Beautyrite collection.. if you buy 2 BeautyBase products you get a free bag! 🥰                                                                                 | 0.000 | 0.688 | 0.312 |     0.8977 |
| 52 | I would like a girls winter date where we shop at Beautyrite. We have tea, read books together &amp; excitingly share thoughts. Then we go home &amp; watch Harry Potter in comfy clothes and blankets. Pls &amp; ty | 0.000 | 0.742 | 0.258 |     0.8591 |
|  5 | My kid has entered the “how many chemicals and colors can I put on my face in one day” phase. Please help, like in the form of Beautyrite gift cards?                                                                | 0.000 | 0.714 | 0.286 |     0.8555 |
| 11 | Linda has 50 items in her Beautyrite cart😃 anyone wanna help a brotha out cuz ima be broke after the holidays😃                                                                                                     | 0.072 | 0.665 | 0.263 |     0.7579 |
| 55 | used my beautychai skincare today that i got via my favorite Beautyrite employee                                                                                                                                     | 0.000 | 0.67  | 0.33  |     0.7506 |
|    | feeling rlly good abt beauty2lips                                                                                                                                                                                    |       |       |       |            |
| 43 | So there's another 20% off Beautyrite... what was the point of the VIP sale?!?!                                                                                                                                      | 0.000 | 0.723 | 0.277 |     0.7164 |
|  2 | Why don’t airports have helpful stores like Beautyrite? Why do we need 5 different burger restaurants?                                                                                                               | 0.000 | 0.712 | 0.288 |     0.6868 |

```python
### BEGIN SOLUTION
df_top10_positive = df_merged.sort_values('compound', ascending=False).head(10)
### END SOLUTION

df_top10_positive
```

Output:
```
text  \
37                                                                                                                                    My BFF gave me a Beautyrite gift card for my birthday 🥰 i love getting gift cards.   
24                                                                                                                                                                                           IN LOVE WITH BEAUTY2LIPS 🤩🥰   
19                                                                                                                                                            Got my BeautyBase merch I am happy I got the first batch 🥰   
6                                                                                    Beautyrite is 20% off today (all purchases). 30% off Beautyrite collection.. if you buy 2 BeautyBase products you get a free bag! 🥰   
52  I would like a girls winter date where we shop at Beautyrite. We have tea, read books together &amp; excitin
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-03, points: 1, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE BELOW
_test_case = 'find-top-10-positive-tweets'
_points = 1
_obfuscate = True

df_top10_positive_SOL = df_merged_SOL.sort_values('compound', ascending=False).head(10)
pd.testing.assert_frame_equal(df_top10_positive.reset_index(drop=True),
                              df_top10_positive_SOL.reset_index(drop=True))
```

---

### 🎯 Challenge 6: Find the top 10 negative tweets

#### 👇 Tasks

- ✔️ Using `df_merged`, filter the top 10 negative tweets.
    - Use the compound score (`compound` column).
- ✔️ Store the filtered DataFrame to a new variable named `df_top10_negative`.
- ✔️ Sort `df_top10_negative` by `compound` in ascending order.

#### 🔑 Expected Output

|    | text                                                                                                                                                                                                              |   neg |   neu |   pos |   compound |
|---:|:------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------------|------:|------:|------:|-----------:|
|  3 | Y’all go to Beautyrite’s IG and look at their awful spring lineup 😭😭😭😭😭😭                                                                                                                                    | 0.484 | 0.516 | 0.000 |    -0.9666 |
| 58 | I feel so bad that my girls Beautyrite package got stolen. She worked so hard for that . . . fucking thieves                                                                                                      | 0.422 | 0.578 | 0.000 |    -0.9209 |
| 59 | 😭😭You can DoorDash Beautyrite now??😭😭                                                                                                                                                                         | 0.515 | 0.485 | 0.000 |    -0.9146 |
| 28 | Ok but the Beautyrite birthday gifts are lowkey a curse cause why am I about to spend $100 for a face cream 😭                                                                                                    | 0.274 | 0.677 | 0.049 |    -0.8519 |
| 46 | I BOUGHT THE WRONG BEAUTYRITE LIPSTICK SHADEEE AND I BASICALLY WASTED MONEY OML IN MY DEFENSE I WAS IN A RUSH THAT LIPSTICK MAKES ME LOOK PALE AF THIS IS WHY I ONLY BUY EYELINER AND MASCARA EVERY SINGLE TIME🙄 | 0.152 | 0.804 | 0.044 |    -0.7603 |
| 35 | I have an interview at Beautyrite today tell my why I had a dream I didn't get the job😂😭😭😭                                                                                                                    | 0.255 | 0.614 | 0.132 |    -0.6597 |
| 14 | why are Beautyrite employees so rude like damn                                                                                                                                                                    | 0.468 | 0.342 | 0.190 |    -0.6189 |
| 13 | cried in the middle of Beautyrite because the workers wouldn’t leave me alone 🤩                                                                                                                                  | 0.345 | 0.655 | 0.000 |    -0.5859 |
| 49 | I hate Beautyrite, I feel out of place                                                                                                                                                                            | 0.346 | 0.654 | 0.000 |    -0.5719 |
|  0 | every time I go into Beautyrite everybody’s makeup look bad, what’s that about?                                                                                                                                   | 0.226 | 0.774 | 0.000 |    -0.5423 |

```python
### BEGIN SOLUTION
df_top10_negative = df_merged.sort_values('compound').head(10)
### END SOLUTION

df_top10_negative
```

Output:
```
text  \
3                                                                                                                                           Y’all go to Beautyrite’s IG and look at their awful spring lineup 😭😭😭😭😭😭   
58                                                                                                      I feel so bad that my girls Beautyrite package got stolen. She worked so hard for that . . . fucking thieves   
59                                                                                                                                                                             😭😭You can DoorDash Beautyrite now??😭😭   
28                                                                                                     Ok but the Beautyrite birthday gifts are lowkey a curse cause why am I about to spend $100 for a face cream 😭   
46  I BOUGHT THE WRONG BEAUTYRITE LIPSTICK SHADEEE AND I BASICALLY WASTED MONEY OML IN MY DEFENSE I WAS IN A RUSH THAT LIPSTICK 
```

#### 🧭 Check Your Work

- Once you're done, run the code cell below to test correctness.
- ✔️ If the code cell runs without an error, you're good to move on.
- ❌ If the code cell throws an error, go back and fix incorrect parts.

<!-- grade_id: challenge-03, points: 1, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE BELOW
_test_case = 'find-top-10-negative-tweets'
_points = 1
_obfuscate = True

df_top10_negative_SOL = df_merged_SOL.sort_values('compound').head(10)
pd.testing.assert_frame_equal(df_top10_negative.reset_index(drop=True),
                              df_top10_negative_SOL.reset_index(drop=True))
```