# Lecture 18 - SQL Queries and Aggregations

Thusrday 2022/03/24

---

## ✨ Aggregating Data

Today, we will review basic SQL queries and try out a few aggregations. We will perform every exercise using both Pandas and SQL to see how similar they are.

▶️ First, run the code cell below to import modules used for **🧭 Check Your Work** sections and the autograder.

```python
import unittest
import base64
tc = unittest.TestCase()
```

---

### 🎯 Pre-exercise: Import Packages

#### 👇 Tasks

- ✔️ Import the following Python packages.
    1. `pandas`: Use alias `pd`.
    2. `numpy`: Use alias `np`.
    3. `sqlite3`: No alias

```python
### BEGIN SOLUTION
import pandas as pd
import numpy as np
import sqlite3
### END SOLUTION
```

#### 🧭 Check your work

<!-- grade_id: pre-exercise, points: 1, grade: true, solution: false -->
```python
import sys
tc.assertTrue('pd' in globals(), 'Check whether you have correctly imported Pandas with an alias.')
tc.assertTrue('np' in globals(), 'Check whether you have correctly imported NumPy with an alias.')
tc.assertTrue('sqlite3' in globals(), 'Check whether you have correctly imported the sqlite3 package.')
```

---
### 📌 Populate a database table from a CSV file

▶️ Run the code below to create a DataFrame named `df_listings`.

```python
df_listings = pd.read_csv('https://github.com/bdi475/datasets/raw/main/boston-airbnb-listings-small.csv')
df_listings_backup = df_listings.copy()

display(df_listings.head(5))
```

Output:
```
name neighbourhood  \
0                          HARBORSIDE-Walk to subway   East Boston   
1     **$49 Special **   Private! Minutes to center!       Roxbury   
2                     $99 Special!! Home Away! Condo       Roxbury   
3                     Bright 1bed facing Golden Dome      Downtown   
4  The Dorset Redline | 3BR 1BA | Walk to Redline...    Dorchester   

         room_type  bedrooms  bathrooms  minimum_nights  price  \
0  Entire home/apt       1.0        1.0              28  150.0   
1  Entire home/apt       1.0        1.0               3  145.0   
2  Entire home/apt       1.0        1.0               3  169.0   
3  Entire home/apt       1.0        1.0              91   81.0   
4  Entire home/apt       3.0        1.0              32  129.0   

   availability_365  number_of_reviews  review_score  is_superhost  
0               123                 17          99.0             0  
1                 0                107          95.0             1  
2                 0    
```

▶️ Run the code below to populate the `listings` table. All data in `df_listings` will be inserted to the table.

```python
conn = sqlite3.connect('airbnb-boston-medium.db')
c = conn.cursor()

tables = list(pd.read_sql_query('SELECT * FROM sqlite_master WHERE type="table";', con=conn)['tbl_name'])

if 'listings' in tables:
    c.execute(f'DELETE FROM listings')
    conn.commit()
    
df_listings.to_sql(name='listings', index=False, con=conn, if_exists='append')

conn.close()
```

#### 🧭 Check your work

```python
# DO NOT CHANGE THE CODE IN THIS CELL
conn_checker = sqlite3.connect('airbnb-boston-medium.db')
table_to_check = 'listings'

# Check if table exists
user_tables = list(pd.read_sql_query('SELECT * FROM sqlite_master WHERE type="table";', con=conn_checker)['tbl_name'])
tc.assertTrue(table_to_check in user_tables, f'{table_to_check} does not exist in your NWT.db file!')

conn_checker.close()
```

▶️ Run the code below to display `df_listings`.

```python
display(df_listings)
```

Output:
```
name neighbourhood  \
0                             HARBORSIDE-Walk to subway   East Boston   
1        **$49 Special **   Private! Minutes to center!       Roxbury   
2                        $99 Special!! Home Away! Condo       Roxbury   
3                        Bright 1bed facing Golden Dome      Downtown   
4     The Dorset Redline | 3BR 1BA | Walk to Redline...    Dorchester   
...                                                 ...           ...   
1334  Quiet and Sunny Top Floor Studio! South End LOVE!     South End   
1335           FOUND Boston Common - Standard Full Room   Bay Village   
1336                   FOUND Boston Common - Queen Room      Downtown   
1337            FOUND Boston Common - Double Queen Room      Downtown   
1338                  FOUND Boston Common - Single Room      Downtown   

            room_type  bedrooms  bathrooms  minimum_nights  price  \
0     Entire home/apt       1.0        1.0              28  150.0   
1     Entire home/apt       1.0     
```

---

### 🎯 Exercise 1: Select all columns and rows from the `listings` table (SQL)

#### 👇 Tasks

- ✔️ Write a query that selects all columns for all rows.
- ✔️ Store your query to a new variable named `query_select_all`.

```python
### BEGIN SOLUTION
query_select_all = '''
SELECT *
FROM listings
'''
### END SOLUTION

conn = sqlite3.connect('airbnb-boston-medium.db')
df_result = pd.read_sql_query(query_select_all, con=conn)
display(df_result)
conn.close()
```

Output:
```
name neighbourhood  \
0                             HARBORSIDE-Walk to subway   East Boston   
1        **$49 Special **   Private! Minutes to center!       Roxbury   
2                        $99 Special!! Home Away! Condo       Roxbury   
3                        Bright 1bed facing Golden Dome      Downtown   
4     The Dorset Redline | 3BR 1BA | Walk to Redline...    Dorchester   
...                                                 ...           ...   
1334  Quiet and Sunny Top Floor Studio! South End LOVE!     South End   
1335           FOUND Boston Common - Standard Full Room   Bay Village   
1336                   FOUND Boston Common - Queen Room      Downtown   
1337            FOUND Boston Common - Double Queen Room      Downtown   
1338                  FOUND Boston Common - Single Room      Downtown   

            room_type  bedrooms  bathrooms  minimum_nights  price  \
0     Entire home/apt       1.0        1.0              28  150.0   
1     Entire home/apt       1.0     
```

#### 🧭 Check your work

<!-- grade_id: exercise-01, points: 1, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE IN THIS CELL
conn = sqlite3.connect('airbnb-boston-medium.db')
decoded_query = base64.b64decode(b'ClNFTEVDVCAqCkZST00gbGlzdGluZ3MK').decode()
df_check = pd.read_sql_query(decoded_query, con=conn)
tc.assertEqual(df_result.shape, df_check.shape, 'Number of rows and/or columns is different')
conn.close()
```

---

### 🎯 Exercise 2A: Find expensive listings (Pandas)

#### 👇 Tasks

- ✔️ Using `df_listings`, find all rows with a price of over $1,000 (`df_listings['price'] > 1000`).
- ✔️ Store the result to a new variable named `df_expensive_listings`.

#### 🔑 Expected Output

Your index column may contain different values.

|     | name                                           | neighbourhood   | room_type       |   bedrooms |   bathrooms |   minimum_nights |   price |   availability_365 |   number_of_reviews |   review_score |   is_superhost |
|----:|:-----------------------------------------------|:----------------|:----------------|-----------:|------------:|-----------------:|--------:|-------------------:|--------------------:|---------------:|---------------:|
| 502 | Large 4BDR near Harvard with Parking           | Allston         | Entire home/apt |          4 |         2.5 |                1 |    1250 |                  0 |                 136 |             98 |              0 |
| 669 | The Historic House in the North End/Waterfront | North End       | Entire home/apt |          3 |         2   |               91 |    1052 |                  0 |                 104 |             98 |              1 |

```python
### BEGIN SOLUTION
df_expensive_listings = df_listings[df_listings['price'] > 1000]
### END SOLUTION

display(df_expensive_listings)
```

Output:
```
name neighbourhood  \
502            Large 4BDR near Harvard with Parking       Allston   
669  The Historic House in the North End/Waterfront     North End   

           room_type  bedrooms  bathrooms  minimum_nights   price  \
502  Entire home/apt       4.0        2.5               1  1250.0   
669  Entire home/apt       3.0        2.0              91  1052.0   

     availability_365  number_of_reviews  review_score  is_superhost  
502                 0                136          98.0             0  
669                 0                104          98.0             1
```

#### 🧭 Check your work

<!-- grade_id: exercise-2A, points: 1, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE IN THIS CELL
df_check = df_listings_backup.query(f"{'PrIcE'.lower()} > {25 * 40}")

pd.testing.assert_frame_equal(df_expensive_listings.sort_values(df_expensive_listings.columns.tolist()).reset_index(drop=True),
                              df_check.sort_values(df_check.columns.tolist()).reset_index(drop=True))
```

---

### 🎯 Exercise 2B: Find expensive listings (SQL)

#### 👇 Tasks

- ✔️ Using the `listings` table, select rows with a price of over $1,000 (`price > 1000`).
- ✔️ Select all columns.
- ✔️ Store your query to a new variable named `query_expensive_listings`.

```python
### BEGIN SOLUTION
query_expensive_listings = '''
SELECT *
FROM listings
WHERE price > 1000
'''
### END SOLUTION

conn = sqlite3.connect('airbnb-boston-medium.db')
df_result = pd.read_sql_query(query_expensive_listings, con=conn)
display(df_result)
conn.close()
```

Output:
```
name neighbourhood  \
0            Large 4BDR near Harvard with Parking       Allston   
1  The Historic House in the North End/Waterfront     North End   

         room_type  bedrooms  bathrooms  minimum_nights   price  \
0  Entire home/apt       4.0        2.5               1  1250.0   
1  Entire home/apt       3.0        2.0              91  1052.0   

   availability_365  number_of_reviews  review_score  is_superhost  
0                 0                136          98.0             0  
1                 0                104          98.0             1
```

#### 🧭 Check your work

<!-- grade_id: exercise-2B, points: 1, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE IN THIS CELL
conn = sqlite3.connect('airbnb-boston-medium.db')
decoded_query = base64.b64decode(b'ClNFTEVDVCAqCkZST00gbGlzdGluZ3MKV0hFUkUgcHJpY2UgPiAxMDAwCg==').decode()
df_check = pd.read_sql_query(decoded_query, con=conn)
tc.assertEqual(df_result.shape, df_check.shape, 'Number of rows and/or columns is different')
pd.testing.assert_frame_equal(df_result.sort_values(df_result.columns.tolist()).reset_index(drop=True),
                              df_check.sort_values(df_check.columns.tolist()).reset_index(drop=True))
conn.close()
```

---

### 🎯 Exercise 3A: Find large listings (Pandas)

#### 👇 Tasks

- ✔️ Using `df_listings`, find all rows where the listing has:
    - 5 or more bedrooms
    - **AND** 3 or more bathrooms
- ✔️ Select only the following 5 columns (in the same order):
    - `name`, `bedrooms`, `bathrooms`, `price`, `review_score`
- ✔️ Store the result to a new variable named `df_large_listings`.

#### 🔑 Expected Output

Your index column may contain different values.

|     | name                                               |   bedrooms |   bathrooms |   price |   review_score |
|----:|:---------------------------------------------------|-----------:|------------:|--------:|---------------:|
| 151 | Beautiful Philadelphia house near Harvard U. w/ pk |          5 |           3 |     600 |             82 |
| 250 | Spacious 4 BR | 2.5 BA Single-Family Home.         |          5 |           3 |     373 |             97 |
| 921 | ★Large Retreat 5BR w/3BA★ Close to Everything      |          5 |           3 |     738 |             95 |

```python
### BEGIN SOLUTION
df_large_listings = df_listings[(df_listings['bedrooms'] >= 5) & (df_listings['bathrooms'] >= 3)]
df_large_listings = df_large_listings[['name', 'bedrooms', 'bathrooms', 'price', 'review_score']]
### END SOLUTION

display(df_large_listings)
```

Output:
```
name  bedrooms  bathrooms  \
151  Beautiful Philadelphia house near Harvard U. w...       5.0        3.0   
250         Spacious 4 BR | 2.5 BA Single-Family Home.       5.0        3.0   
921      ★Large Retreat 5BR w/3BA★ Close to Everything       5.0        3.0   

     price  review_score  
151  600.0          82.0  
250  373.0          97.0  
921  738.0          95.0
```

#### 🧭 Check your work

<!-- grade_id: exercise-3A, points: 1, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE IN THIS CELL
__r = 'oo'.join(['R', 'Ms']).lower()
df_check = df_listings.query(f'bed{__r} > {2 ** 2} & bath{__r} > {2 ** 1}') \
    [['_'.join(['review', 'score']), 'price', 'bath' + __r, 'bed' + __r, 'name'][::-1]]

pd.testing.assert_frame_equal(df_large_listings.sort_values(df_large_listings.columns.tolist()).reset_index(drop=True),
                              df_check.sort_values(df_check.columns.tolist()).reset_index(drop=True))
```

---

### 🎯 Exercise 3B: Find large listings (SQL)

#### 👇 Tasks

- ✔️ Using the `listings` table, find all rows where the listing has:
    - 5 or more bedrooms
    - **AND** 3 or more bathrooms
- ✔️ Select only the following 5 columns (in the same order):
    - `name`, `bedrooms`, `bathrooms`, `price`, `review_score`
- ✔️ Store your query to a new variable named `query_large_listings`.

```python
### BEGIN SOLUTION
query_large_listings = '''
SELECT name, bedrooms, bathrooms, price, review_score
FROM listings
WHERE (bedrooms >= 5) AND (bathrooms >= 3)
'''
### END SOLUTION

conn = sqlite3.connect('airbnb-boston-medium.db')
df_result = pd.read_sql_query(query_large_listings, con=conn)
display(df_result)
conn.close()
```

Output:
```
name  bedrooms  bathrooms  \
0  Beautiful Philadelphia house near Harvard U. w...       5.0        3.0   
1         Spacious 4 BR | 2.5 BA Single-Family Home.       5.0        3.0   
2      ★Large Retreat 5BR w/3BA★ Close to Everything       5.0        3.0   

   price  review_score  
0  600.0          82.0  
1  373.0          97.0  
2  738.0          95.0
```

#### 🧭 Check your work

<!-- grade_id: exercise-3B, points: 1, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE IN THIS CELL
conn = sqlite3.connect('airbnb-boston-medium.db')
decoded_query = base64.b64decode(b'ClNFTEVDVCBuYW1lLCBiZWRyb29tcywgYmF0aHJvb21zLCBwcmljZSwgcmV2aWV3X3Njb3JlCkZST00gbGlzdGluZ3MKV0hFUkUgKGJlZHJvb21zID49IDUpIEFORCAoYmF0aHJvb21zID49IDMpCg=='
).decode()
df_check = pd.read_sql_query(decoded_query, con=conn)
tc.assertEqual(df_result.columns.tolist(), df_check.columns.tolist(), 'Incorrect set of columns or order')
tc.assertEqual(df_result.shape, df_check.shape, 'Number of rows and/or columns is different')
pd.testing.assert_frame_equal(df_result.sort_values(df_result.columns.tolist()).reset_index(drop=True),
                              df_check.sort_values(df_check.columns.tolist()).reset_index(drop=True))
conn.close()
```

---

### 🎯 Exercise 4A: Number of listings (Pandas)

#### 👇 Tasks

- ✔️ Find the number of rows in `df_listings`.
- ✔️ Store the result to a new variable named `num_listings`.
- ✔️ `num_listings` should be an integer type.

```python
### BEGIN SOLUTION
num_listings = df_listings.shape[0]
### END SOLUTION

print(f'There are {num_listings} AirBnB listings in the dataset.')
```

Output:
```
There are 1339 AirBnB listings in the dataset.
```

#### 🧭 Check your work

<!-- grade_id: exercise-4A, points: 1, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE IN THIS CELL
tc.assertEqual(num_listings, len(df_listings_backup.index), f'Incorrect number of listings - should be {len(df_listings_backup.index)}')
```

---

### 🎯 Exercise 4B: Number of listings (SQL)

#### 👇 Tasks

- ✔️ Write a query that counts the number of rows in the `listings` table.
- ✔️ Store your query in a new variable named `query_num_listings`.

```python
### BEGIN SOLUTION
query_num_listings = '''
SELECT COUNT(*)
FROM listings
'''
### END SOLUTION

conn = sqlite3.connect('airbnb-boston-medium.db')
df_result = pd.read_sql_query(query_num_listings, con=conn)
display(df_result)
conn.close()
```

Output:
```
COUNT(*)
0      1339
```

#### 🧭 Check your work

<!-- grade_id: exercise-4B, points: 1, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE IN THIS CELL
conn = sqlite3.connect('airbnb-boston-medium.db')
decoded_query = base64.b64decode(b'ClNFTEVDVCBDT1VOVCgqKQpGUk9NIGxpc3RpbmdzCg==').decode()
df_check = pd.read_sql_query(decoded_query, con=conn)
tc.assertEqual(df_result.iloc[0, 0], df_check.iloc[0, 0], 'Incorrect number of listings')
conn.close()
```

---

### 🎯 Exercise 5A: Get the number of listings by neighbourhood (Pandas)

#### 👇 Tasks

- ✔️ Using `df_listings`, find the number of listings by neighbourhood.
- ✔️ Store the result to a new variable named `df_by_neighbourhood`.
- ✔️ `df_by_neighbourhood` should have the following two columns.
    1. `neighbourhood`: Name of the neighbourhood
    2. `num_listings`: Number of listings in a neighbourhood
- ✔️ `print(df_by_neighbourhood.columns.tolist())` should print `['neighbourhood', 'num_listings']`.
- ✔️ **Sort** `df_by_neighbourhood` by `num_listings` in descending order.

#### 🔑 Expected Output

Your index column may contain different values.

|    | neighbourhood   |   num_listings |
|---:|:----------------|---------------:|
|  7 | Dorchester      |            215 |
| 12 | Jamaica Plain   |            129 |
| 19 | Roxbury         |            107 |
|  9 | East Boston     |             90 |
| 22 | South End       |             86 |

```python
### BEGIN SOLUTION
df_by_neighbourhood = df_listings.groupby('neighbourhood', as_index=False).agg({
    'name': 'count'
})

df_by_neighbourhood.rename(columns={'name': 'num_listings'}, inplace=True)
df_by_neighbourhood.sort_values('num_listings', ascending=False, inplace=True)
### END SOLUTION

df_by_neighbourhood.head(5)
```

Output:
```
neighbourhood  num_listings
7      Dorchester           215
12  Jamaica Plain           129
19        Roxbury           107
9     East Boston            90
22      South End            86
```

#### 🧭 Check your work

<!-- grade_id: exercise-5A, points: 1, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE IN THIS CELL
df_check = df_listings_backup.pivot_table(index=['neigh' + 'BoUr'.lower() + 'HooD'.lower()],
                                          values=['NaMe'.lower()],
                                          aggfunc=('CoU' + 'Nt').lower()) \
                                .rename(columns={'NaME'.lower(): '_'.join(['NuM', 'LiSTinGs']).lower()}) \
                                .reset_index().sort_values('_'.join(['NuM', 'LiSTinGs']).lower(), ascending=bool(0))

pd.testing.assert_series_equal(df_by_neighbourhood['num_listings'].reset_index(drop=True),
                               df_check['num_listings'].reset_index(drop=True))
```

---

### 🎯 Exercise 5B: Get the number of listings by neighbourhood (SQL)

#### 👇 Tasks

- ✔️ Write a query that counts the number of listings by neighbourhood from the `listings` table.
- ✔️ The result of the query should have the following two columns.
    1. `neighbourhood`: Name of the neighbourhood
    2. `num_listings`: Number of listings in a neighbourhood
- ✔️ Store your query in a new variable named `query_num_listings_by_neighbourhood`.

```python
### BEGIN SOLUTION
query_num_listings_by_neighbourhood = '''
SELECT neighbourhood, COUNT(*) AS num_listings
FROM listings
GROUP BY neighbourhood
ORDER BY num_listings DESC
'''
### END SOLUTION

conn = sqlite3.connect('airbnb-boston-medium.db')
df_result = pd.read_sql_query(query_num_listings_by_neighbourhood, con=conn)
display(df_result)
conn.close()
```

Output:
```
neighbourhood  num_listings
0                Dorchester           215
1             Jamaica Plain           129
2                   Roxbury           107
3               East Boston            90
4                 South End            86
5                  Downtown            82
6                   Allston            82
7              South Boston            67
8                  Brighton            66
9               Beacon Hill            54
10                 Back Bay            50
11              Charlestown            47
12                   Fenway            42
13               Roslindale            38
14                North End            33
15                 Mattapan            30
16                Hyde Park            25
17             West Roxbury            24
18             Mission Hill            23
19              Bay Village            22
20                 West End            12
21                Chinatown             7
22  South Boston Waterfront             6
23    
```

#### 🧭 Check your work

<!-- grade_id: exercise-5B, points: 1, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE IN THIS CELL
conn = sqlite3.connect('airbnb-boston-medium.db')
decoded_query = base64.b64decode(b'ClNFTEVDVCBuZWlnaGJvdXJob29kLCBDT1VOVCgqKSBBUyBudW1fbGlzdGluZ3MKRlJPTSBsaXN0aW5ncwpHUk9VUCBCWSBuZWlnaGJvdXJob29kCk9SREVSIEJZIG51bV9saXN0aW5ncyBERVNDCg==').decode()
df_check = pd.read_sql_query(decoded_query, con=conn)
tc.assertEqual(df_result.shape, df_check.shape, 'Number of rows and/or columns is different')
pd.testing.assert_series_equal(df_result['num_listings'].reset_index(drop=True),
                               df_result['num_listings'].reset_index(drop=True))
conn.close()
```

#### 🚀 Solution

```sql
SELECT neighbourhood, COUNT(*) AS num_listings
FROM listings
GROUP BY neighbourhood
ORDER BY num_listings DESC
```

---

### 🎯 Exercise 6A: Average price by room type (Pandas)

#### 👇 Tasks

- ✔️ Using `df_listings`, find the average price by room type.
- ✔️ Store the result to a new variable named `df_price_by_room_type`.
- ✔️ `df_price_by_room_type` should have the following two columns.
    1. `room_type`: Room type
    2. `price`: Average price of the room type
- ✔️ `print(df_price_by_room_type.columns.tolist())` should print `['room_type', 'price']`.
- ✔️ **Sort** `df_price_by_room_type` by `price` in ascending order.

#### 🔑 Expected Output

Your index column may contain different values.

|    | room_type       |    price |
|---:|:----------------|---------:|
|  3 | Shared room     |  34.75   |
|  2 | Private room    |  82.2278 |
|  0 | Entire home/apt | 201.135  |
|  1 | Hotel room      | 206.273  |

```python
### BEGIN SOLUTION
df_price_by_room_type = df_listings.groupby('room_type', as_index=False).agg({
    'price': 'mean'
})

df_price_by_room_type.sort_values('price', inplace=True)

### END SOLUTION

df_price_by_room_type
```

Output:
```
room_type       price
3      Shared room   34.750000
2     Private room   82.227848
0  Entire home/apt  201.134890
1       Hotel room  206.272727
```

#### 🧭 Check your work

<!-- grade_id: exercise-6A, points: 1, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE IN THIS CELL
df_check = df_listings_backup.pivot_table(index=['RoOm_'.lower() + 'TyPe'.lower()],
                                          values=['price'],
                                          aggfunc=np.mean) \
                                .reset_index() \
                                .sort_values('price', ascending=bool(not False))

pd.testing.assert_frame_equal(df_price_by_room_type.sort_values(df_price_by_room_type.columns.tolist()).reset_index(drop=True),
                              df_check.sort_values(df_check.columns.tolist()).reset_index(drop=True))
```

---

### 🎯 Exercise 6B: Average price by room type (SQL)

#### 👇 Tasks

- ✔️ Write a query that finds the average price by room type from the `listings` table.
- ✔️ Store your query to a new variable named `query_price_by_room_type`.
- ✔️ Your result should have the following two columns.
    1. `room_type`: Room type
    2. `price`: Average price of the room type
- ✔️ **Sort** your result by `price` in ascending order.

```python
### BEGIN SOLUTION
query_price_by_room_type = '''
SELECT room_type, AVG(price) as price
FROM listings
GROUP BY room_type
ORDER BY price
'''
### END SOLUTION

conn = sqlite3.connect('airbnb-boston-medium.db')
df_result = pd.read_sql_query(query_price_by_room_type, con=conn)
display(df_result)
conn.close()
```

Output:
```
room_type       price
0      Shared room   34.750000
1     Private room   82.227848
2  Entire home/apt  201.134890
3       Hotel room  206.272727
```

#### 🧭 Check your work

<!-- grade_id: exercise-6B, points: 1, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE IN THIS CELL
conn = sqlite3.connect('airbnb-boston-medium.db')
decoded_query = base64.b64decode(b'ClNFTEVDVCByb29tX3R5cGUsIEFWRyhwcmljZSkgYXMgcHJpY2UKRlJPTSBsaXN0aW5ncwpHUk9VUCBCWSByb29tX3R5cGUKT1JERVIgQlkgcHJpY2UK').decode()
df_check = pd.read_sql_query(decoded_query, con=conn)
tc.assertEqual(df_result.shape, df_check.shape, 'Number of rows and/or columns is different')
pd.testing.assert_frame_equal(df_result.sort_values(df_result.columns.tolist()).reset_index(drop=True),
                              df_check.sort_values(df_check.columns.tolist()).reset_index(drop=True))

conn.close()
```