# Lecture 19 - SQL Joins

Tuesday 2022/03/29

---

## ✨ Joining Tables

Today, we will review basic SQL joins.

▶️ First, run the code cell below to import modules used for **🧭 Check Your Work** sections and the autograder.

```python
import unittest
import base64
tc = unittest.TestCase()
```

---

### 🎯 Pre-exercise: Import Packages

#### 👇 Tasks

- ✔️ Import the following Python packages.
    1. `pandas`: Use alias `pd`.
    2. `numpy`: Use alias `np`.
    3. `sqlite3`: No alias

```python
### BEGIN SOLUTION
import pandas as pd
import numpy as np
import sqlite3
### END SOLUTION
```

#### 🧭 Check your work

<!-- grade_id: pre-exercise, points: 1, grade: true, solution: false -->
```python
import sys
tc.assertTrue('pd' in globals(), 'Check whether you have correctly imported Pandas with an alias.')
tc.assertTrue('np' in globals(), 'Check whether you have correctly imported NumPy with an alias.')
tc.assertTrue('sqlite3' in globals(), 'Check whether you have correctly imported the sqlite3 package.')
```

---
### 📌 Read Sqlite Database File

▶️ Run the code below to select the first 5 rows from the `students`, `courses`, and `enrollments` tables.

```python
# DO NOT CHANGE THE CODE IN THIS CELL
conn = sqlite3.connect('course-enrollments-sample.db')

print('students table')
display(pd.read_sql_query('SELECT * FROM students LIMIT 5;', con=conn))

print('========================')

print('courses table')
display(pd.read_sql_query('SELECT * FROM courses LIMIT 5;', con=conn))

print('========================')

print('enrollments table')
display(pd.read_sql_query('SELECT * FROM enrollments LIMIT 5;', con=conn))

print('========================')

conn.close()
```

Output:
```
students table
   student_id first_name   last_name        major      class
0       69540      Megan  Cunningham      History  Sophomore
1       60083      Wendy       House   Statistics     Senior
2       65353    Anthony        Hill    Sociology   Freshman
3       65070    Kenneth      Hayden  Accountancy     Junior
4       63215     Willie    Anderson      History   Freshman
========================
courses table
  course_id                       course_name  credit_hours
0  ATMS 120      Severe and Hazardous Weather             3
1    CS 125  Introduction to Computer Science             4
2  PHYS 140                   How Things Work             4
3  CLCV 115      Mythology of Greece and Rome             3
4  ASTR 122                Stars and Galaxies             2
========================
enrollments table
   student_id course_id
0       63438  ASTR 122
1       76342    CS 125
2       64798  PHYS 140
3       64932  ECON 101
4       60489  PHYS 140
========================
```

#### 🧭 Check your work

```python
# DO NOT CHANGE THE CODE IN THIS CELL
conn_checker = sqlite3.connect('course-enrollments-sample.db')
tables_to_check = ['students', 'courses', 'enrollments']

# Check if table exists
user_tables = list(pd.read_sql_query('SELECT * FROM sqlite_master WHERE type="table";', con=conn_checker)['tbl_name'])

for table_to_check in tables_to_check:
    tc.assertTrue(table_to_check in user_tables, f'{table_to_check} does not exist in your NWT.db file!')

conn_checker.close()
```

---

### 🎯 Exercise 1: Join `students` into `enrollments` table

#### 👇 Tasks

- ✔️ Write a query that joins the `students` table into `enrollments`.
- ✔️ Select all columns.
- ✔️ Store your query to a new variable named `query_joined1`.

```python
### BEGIN SOLUTION
query_joined1 = '''
SELECT *
FROM enrollments
LEFT JOIN students
ON enrollments.student_id == students.student_id;
'''
### END SOLUTION

conn = sqlite3.connect('course-enrollments-sample.db')
df_result = pd.read_sql_query(query_joined1, con=conn)
display(df_result)
conn.close()
```

Output:
```
student_id course_id  student_id first_name   last_name  \
0         63438  ASTR 122       63438    Kaitlyn      Butler   
1         76342    CS 125       76342     Joanna     Jackson   
2         64798  PHYS 140       64798     Monica      Reilly   
3         64932  ECON 101       64932       Noah      Nguyen   
4         60489  PHYS 140       60489   Jennifer  Williamson   
..          ...       ...         ...        ...         ...   
197       74780  ATMS 120       74780     Justin      Patton   
198       72793    IB 105       72793      Brian       Smith   
199       66953    CS 125       66953       Rose        Shea   
200       74413  PHYS 140       74413    Bethany      Howard   
201       76695  ECON 101       76695      Megan      Foster   

                major      class  
0    Computer Science     Senior  
1           Sociology     Junior  
2             History  Sophomore  
3           Sociology   Freshman  
4    Computer Science  Sophomore  
..                ...     
```

#### 🧭 Check your work

<!-- grade_id: exercise-01, points: 3, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE IN THIS CELL
conn = sqlite3.connect('course-enrollments-sample.db')
df_check = pd.read_sql_query(query_joined1, con=conn)
tc.assertEqual(df_result.shape, (202, 7), 'Incorrect number of rows and/or columns')
conn.close()
```

---

### 🎯 Exercise 2: Join all three tables

#### 👇 Tasks

- ✔️ Write a query that joins all three tables.
- ✔️ Select all columns.
- ✔️ Store your query to a new variable named `query_joined2`.

```python
### BEGIN SOLUTION
query_joined2 = '''
SELECT *
FROM enrollments
LEFT JOIN students
ON enrollments.student_id == students.student_id
LEFT JOIN courses
ON enrollments.course_id == courses.course_id;
'''
### END SOLUTION

conn = sqlite3.connect('course-enrollments-sample.db')
df_result = pd.read_sql_query(query_joined2, con=conn)
display(df_result)
conn.close()
```

Output:
```
student_id course_id  student_id first_name   last_name  \
0         63438  ASTR 122       63438    Kaitlyn      Butler   
1         76342    CS 125       76342     Joanna     Jackson   
2         64798  PHYS 140       64798     Monica      Reilly   
3         64932  ECON 101       64932       Noah      Nguyen   
4         60489  PHYS 140       60489   Jennifer  Williamson   
..          ...       ...         ...        ...         ...   
197       74780  ATMS 120       74780     Justin      Patton   
198       72793    IB 105       72793      Brian       Smith   
199       66953    CS 125       66953       Rose        Shea   
200       74413  PHYS 140       74413    Bethany      Howard   
201       76695  ECON 101       76695      Megan      Foster   

                major      class course_id                       course_name  \
0    Computer Science     Senior  ASTR 122                Stars and Galaxies   
1           Sociology     Junior    CS 125  Introduction to Computer Science
```

#### 🧭 Check your work

<!-- grade_id: exercise-02, points: 3, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE IN THIS CELL
conn = sqlite3.connect('course-enrollments-sample.db')
df_check = pd.read_sql_query(query_joined2, con=conn)
tc.assertEqual(df_result.shape[0], 202, 'Incorrect number of rows')
conn.close()
```

---

### 🎯 Exercise 3: Total number of credit hours by class standing

#### 👇 Tasks

- ✔️ Write a query that finds the total number of credit hours by class status (freshman, sophomore, junior, senior).
- ✔️ Store your query to a new variable named `query_joined3`.

#### 🔑 Expected Output

|    | class     |   SUM(credit_hours) |
|---:|:----------|--------------------:|
|  0 | Freshman  |                 198 |
|  1 | Junior    |                 137 |
|  2 | Senior    |                 128 |
|  3 | Sophomore |                 199 |

```python
### BEGIN SOLUTION
query_joined3 = '''
SELECT class, SUM(credit_hours)
FROM enrollments
LEFT JOIN students
ON enrollments.student_id == students.student_id
LEFT JOIN courses
ON enrollments.course_id == courses.course_id
GROUP BY class;
'''
### END SOLUTION

conn = sqlite3.connect('course-enrollments-sample.db')
df_result = pd.read_sql_query(query_joined3, con=conn)
display(df_result)
conn.close()
```

Output:
```
class  SUM(credit_hours)
0   Freshman                198
1     Junior                137
2     Senior                128
3  Sophomore                199
```

#### 🧭 Check your work

<!-- grade_id: exercise-03, points: 4, grade: true, solution: false -->
```python
# DO NOT CHANGE THE CODE IN THIS CELL
conn = sqlite3.connect('course-enrollments-sample.db')
df_check = pd.read_sql_query(query_joined3, con=conn)

df_check.columns = ['A', 'B']
df_correct = pd.DataFrame({'A': ['Freshman', 'Junior', 'Senior', 'Sophomore'],
 'B': [198, 137, 128, 199]})

pd.testing.assert_frame_equal(df_check, df_correct)

conn.close()
```